Basolateral amygdala oscillations enable fear learning in a biophysical model

Author response:

The following is the authors’ response to the current reviews.

We thank the Reviewer for all their effort and suggestions over multiple drafts. Their comments have encouraged us to read and think more deeply about the issue under discussion (BLA spiking in response to CS/US inputs), and to find the papers whose contents we think provide a potential solution. We agree that there is more to understand about the mechanisms underlying associative learning in the BLA. We offer our paper as providing a new way of understanding the role of circuit dynamics (rhythms) in guiding associative learning via STDP. As we pointed out in our response to the previous review, the issue highlighted by the Reviewer is an issue for the entire field of associative learning in BLA: our discussion of the issue suggests why the experimentally observed BLA spiking in response to CS inputs, performed in the absence of US inputs (as done in the papers cited by the Reviewer), may not be what occurs in the presence of the US. Since our explanation involves the role of neuromodulators, such as ACh and dopamine, the suggestion is open to further testing.

The following is the authors’ response to the original reviews.

Reviewer #1:

Public Review’s only objection: “Deficient in this study is the construction of the afferent drive to the network, which does elicit activities that are consistent with those observed to similar stimuli. It still remains to be demonstrated that their mechanism promotes plasticity for training protocols that emulate the kinds of activities observed in the BLA during fear conditioning.”

Recommendations for the Authors: “The authors have successfully addressed most of my concerns. I commend them for their thorough response. The one nagging issue is the unrealistic activation used to drive CS and US activation in their network. While I agree that their stimulus parameters are consistent with a contextual fear task, or one that uses an olfactory CS, this was not the focus of their study as originally conceived. Moreover, the types of activation observed in response to auditory cues, which is the focus of their study, do not follow what is reported experimentally. Thus, I stand by the critique that the proposed mechanism has not been demonstrated to work for the conditioning task which the authors sought to emulate (Krabbe et al. 2019). Frustratingly, addressing this is simple: run the model with ECS neurons driven so that they fire bursts of action potentials every ~1 sec for 30 sec, and with the US activation noncontiguous with that. If the model does not produce plasticity in this case, then it suggests that the mechanisms embedded in the model are not sufficient, and more work is needed to identify them. While 'memory' effects are possible that could extend the temporal contiguity of the CS and US, the authors need to provide experimental evidence for this occurring in the BLA under similar conditions if they want to invoke it in their model.

(1) Fair response. I accept the authors arguments and changes.

(2) The authors rightly point out that the simulated afferents need not perfectly match the time courses of the peripheral inputs, since what the amygdala receives them indirectly via the thalamus, cortex, etc. However, it is known how amygdala neurons respond to such stimuli, so it behooves the authors to incorporate that fact into their model.

Quirk et al. 1997 show that the response to the tone plummets after the first 100 ms in Figs 5A and 6B. The Herry et al. 2007 paper emphasizes the transient response to tone pips, with spiking falling back to a poisson low firing rate baseline outside of the time when the pip is delivered.

Regarding potential metabotropic glutamate activation, the stimulus in Whittington et al. 1995 was electrical stimulation at 100 Hz that would synchronously activate a large volume of tissue, which is far outside the physiological norm. I appreciate that metabotropic glutamate receptors may play a role here, but ultimately the model depends upon spiking activity for the plastic process to occur, and to the best of my knowledge the spiking activity in BLA in response to a sustained, unconditioned tone, is brief (see also Quirk, Repa, and Ledoux 1995). Perhaps a better justification for the authors would be Bordi and Ledoux 1992, which found that 18% of auditory responsive neurons showed a 'sustained' response, but the sustained response neurons appear to show much weaker responses than those with transient ones (Fig 2). I am willing to say that their paper IS relevant to contextual fear, but that is not what the authors set out to do.

(3) Fair response.

(4) Very good response!

Minor points: All points were addressed.”

We thank Reviewer 1 (R1) for the positive feedback and also for pointing out that, in R1’s opinion, there is still a nagging issue related to the activation in response to CS we modeled. In (Krabbe et al., 2019), CS is a pulsed input and US is delivered right after the CS offset. The current objection of R1 is that instead, we are modeling CS and US as continuous and overlapping. R1 suggested that we add the actual input and see if they will produce the desired outputs. The answer is simple: it will not work because we need the effects of CS and US on pyramidal cells to overlap. We note that the fear learning community appears to agree with us that such contingency is necessary for synaptic plasticity (Sun et al., 2020; Palchaudhuri et al., 2024). To the best of our understanding, the source of that overlap is not understood in the community, and the gap has been much noticed (Sun et al., 2020). We do note, however, that STDP may not be the only kind of plasticity in fear learning (Li et al., 2009; Kim et al., 2013, 2016).

It is important to emphasize that it is not the aim of our paper to model the origin of the overlap. Rather, our intent is to demonstrate the roles of brain rhythms in producing the appropriate timing for STDP, assuming that ECS and F cells can continue to be active after the offset of CS and US, respectively. This assumption is very close to how the field now treats the plasticity, even for auditory fear conditioning (Sun et al., 2020). Thus, our methodology does not contradict known results. However, the question raised by R1 is indeed very interesting, if not the point of our paper. Hence, below we give details about why our hypothesis is reasonable.

Several papers (Quirk, Repa and LeDoux, 1995; Herry et al, 2007; Bordi and Ledoux 1992) show that the pips in auditory fear conditioning increase the activity of some BLA neurons: after an initial transient, the overall spike rate is still higher than baseline activity. As R1 points out, we did not model the transient increase in BLA spiking activity that occurs in response to each pip in the auditory fear conditioning paradigm. However, we did model the low-level sustained activity that occurs in between pips of the CS in the absence of US (Quirk, Repa and LeDoux, 1995, Fig. 2) and after CS offset (see Fig. 2B, left hand part of our manuscript). We read the data of Quirk et al., 1995 as suggesting that the low-level activity can be sustained for some indefinite time after a pip (cut off of recording was at 500 ms with no noticeable decrease in activity). As such, even if the pips and the US do not overlap in time, as in (Krabbe et al., 2019), the spiking of the ECS can be sustained after CS offset and thus overlap with US, a condition necessary in our model for plasticity through STDP. In Herry et al., 2007 Fig. 3 shows that BLA neurons respond to a pip at the population level with a transient increase in spiking and return to a baseline Poisson firing rate. However, a subset of cells continues to fire at an increased-over-baseline rate after the transient effect wears off (Fig. 3C, top few neurons) and this increased rate extends to the end of the recording time (here ~ 300 ms). These are the cells we consider to be ECS in our model. In Quirk et al., 1997, Fig. 5A also shows sustained low level activity of neurons in BLA in response to a pip. The low-level activity is shown to increase after fear learning, as is also the case in our model since ECS now entrains F so that there are more pyramidal cells spiking in response to CS. The question remains as to whether the spiking is sustained long enough and at a high enough rate for STDP to take place when US is presented sometime after the stop of the CS.

Experimental recordings cannot speak to the rate of spiking of BLA neurons during US due to recording interference from the shock. However, evidence seems to suggest that ECS activity should increase during the US due to the release of acetylcholine (ACh) from neurons in the basal forebrain (BF) (Rajebhosale et al., 2024). Pyramidal cells of the BLA robustly express M1 muscarinic ACh receptors (Muller et al., 2013; McDonald and Mott, 2021). Thus, ACh from BF should elicit a depolarization in pyramidal cells. Indeed, the pairing of ACh with even low levels of spiking of BLA neurons results in a membrane depolarization that can last 7 – 10 s (Unal et al., 2015). This should induce higher spiking rates and more sustained activity in the ECS and F neurons during and after the presentation of US, thus ensuring a concomitant activation of ECS and fear (F) neurons necessary for STDP to take place. Other modulators, including dopamine, may also play a role in producing the sustained activity. Activation of US leads to increased dopamine release in the BLA (Harmer and Phillips, 1999; Suzuki et al., 2002). D1 receptors are known to increase the membrane excitability of BLA projection neurons by lowering their spiking threshold (Kröner et al., 2005). Thus, the activation of the US can lead to continued and higher firing rates of ECS and F. The effect of dopamine can last up to 20 minutes (Kröner et al., 2005). For CS-positive neurons, the ACh modulation coming from the firing of US may lead to a temporary extension of firing that is then amplified and continued by dopaminergic effects.

Hence, we suggest that a solution to the problem raised by R1 may be solved by considering the roles of ACh and dopamine in the BLA. The involvement of neuromodulators is consistent with the suggestion of (Sun et al., 2020). The model we have may be considered a “minimal” model that puts in by hand the overlap in activity due to the neuromodulation without explicitly modeling it. As R1 says, it is important for us to give the motivation of our hypotheses. We have used the simplest way to model overlap without assumptions about timing specificity in the overlap.

To account for these points in the manuscript, we first specified that we consider the effects of the US and CS inputs on the neuronal network as overlapping, while the actual inputs may not overlap. To do that, we added the following text:

(1) In the introduction:

“In this paper, we aim to show 1) How a variety of BLA interneurons (PV, SOM and VIP) lead to the creation of these rhythms and 2) How the interaction of the interneurons and the rhythms leads to the appropriate timing of the cells responding to the US and those responding to the CS to promote fear association through spike-timing-dependent plasticity (STDP). Since STDP requires overlap of the effects of the CS and US, and some conditioning paradigms do not have overlapping US and CS, we include as a hypothesis that the effects of the CS and US overlap even if the CS and US stimuli do not. In the Discussion, we suggest how neuromodulation by ACh and/or dopamine can provide such overlap. We create a biophysically detailed model of the BLA circuit involving all three types of interneurons and show how each may participate in producing the experimentally observed rhythms and interacting to produce the necessary timing for the fear learning.”

(2) In the Result section “With the depression-dominated plasticity rule, all interneuron types are needed to provide potentiation during fear learning”:

“The 40-second interval we consider has both ECS and F, as well as VIP and PV interneurons, active during the entire period: an initial bout of US is known to produce a long-lasting fear response beyond the offset of the US (Hole and Lorens, 1975) and to induce the release of neuromodulators. The latter, in particular acetylcholine and dopamine that are known to be released upon US presentation (Harmer and Phillips, 1999; Suzuki et al., 2002; Rajebhosale et al., 2024), may induce more sustained activity in the ECS, F, VIP, and PV neurons during and after the presentation of US, thus ensuring a concomitant activation of those neurons necessary for STDP to take place (see “Assumptions and predictions of the model” in the Discussion).”

(3) In the Discussion section “Synaptic plasticity in our model”:

“Synaptic plasticity is the mechanism underlying the association between neurons that respond to the neutral stimulus CS (ECS) and those that respond to fear (F), which instantiates the acquisition and expression of fear behavior. One form of experimentally observed long-term synaptic plasticity is spike-timing-dependent plasticity (STDP), which defines the amount of potentiation and depression for each pair of pre- and postsynaptic neuron spikes as a function of their relative timing (Bi and Poo, 2001; Caporale and Dan, 2008). All forms of STDP require that there be an overlap in the firing of the pre- and postsynaptic cells. In some fear learning paradigms, the US and the CS do not overlap. We address this below under “Assumptions and predictions of the model”, showing how the effects of US and CS on the spiking of the relevant neurons can overlap even in the absence of overlap of US and CS.”

To fully present our reasoning about the origin of the overlap of the effects of US and CS, we modified and added to the last paragraph of the Discussion section “Assumptions and predictions of the model”, which now reads as follows:

“Finally, our model requires the effect of the CS and US inputs on the BLA neuron activity to overlap in time in order to instantiate fear learning through STDP. Such a hypothesis, that learning uses spike-timing-dependent plasticity, is common in the modeling literature (Bi and Poo, 2001; Caporale and Dan, 2008; Markram et al., 2011). Current paradigms of fear conditioning include examples in which the CS and US stimuli do not overlap (Krabbe et al., 2019). Such a condition might seem to rule out the mechanisms in our paper. Nevertheless, the argument below suggests that the effects of the CS and US can cause an overlap in neuronal spiking of ECS, F, VIP, and SOM, even when CS and US inputs do not overlap.

Experimental recordings cannot speak to the rate of spiking of BLA neurons during US due to recording interference from the shock. However, evidence suggests that ECS activity should increase during the US due to the release of acetylcholine (ACh) from neurons in the basal forebrain (BF) (Rajebhosale et al., 2024). Pyramidal cells of the BLA robustly express M1 muscarinic ACh receptors (McDonald and Mott, 2021). Thus, ACh from BF should elicit a depolarization in pyramidal cells. Indeed, the pairing of ACh with even low levels of spiking of BLA neurons results in a membrane depolarization that can last 7 – 10 s (Unal et al., 2015). Other modulators, including dopamine, may also play a role in producing the sustained activity. Activation of US leads to increased dopamine release in the BLA (Harmer and Phillips, 1999; Suzuki et al., 2002). D1 receptors are known to increase the membrane excitability of BLA projection neurons by lowering their spiking threshold (Kröner et al., 2005). Thus, neuromodulator release should induce higher spiking rates and more sustained activity in the ECS and F neurons during and after the presentation of US, thus ensuring a concomitant activation of ECS and fear (F) neurons necessary for STDP to take place. Thus, the activation of the US can lead to continued and higher firing rates of ECS and F. The effect of dopamine can last up to 20 minutes (Kröner et al., 2005). For CS-positive neurons, the ACh modulation coming from the firing of US may lead to a temporary extension of firing that is then amplified and continued by dopaminergic effects.

Hence, we suggest that a solution to the problem apparently posed by the non-overlap US and CS in some paradigms of auditory fear conditioning (Krabbe et al., 2019) may be solved by considering the roles of ACh and dopamine in the BLA. The model we have may be considered a “minimal” model that puts in by hand the overlap in activity due to the neuromodulation without explicitly modeling it. We have used the simplest way to model overlap without assumptions about timing specificity in the overlap. We note that, even though ECS and F neurons have the ability to fire continuously when ACh and dopamine are involved, the participation of the interneurons enforces periodic silence needed for the depression-dominated STDP.”

In the Discussion (in section “Involvement of other brain structures”), we also acknowledged that the overlap between the effects of US and CS in the BLA may be provided by other brain structures by writing the following:

“In our model, the excitatory projection neurons and VIP and PV interneurons show sustained activity during and after the US presentation, thus allowing potentiation through STDP to take place. The medial prefrontal cortex and/or the hippocampus may provide the substrates for the continued firing of the BLA neurons after the 2-second US stimulation. We also discuss below that this network sustained activity may originate from neuromodulator release induced by US (see section “Assumptions and predictions of the model” in the Discussion).”

We also improved our discussion about the (Grewe et al., 2017) paper, which questions Hebbian plasticity in the context of fear conditioning based on several critiques. We included a new section in the Discussion entitled “Is STDP needed in fear conditioning?” to discuss those critiques and how our model may address them, which reads as follows:

“Is STDP needed in fear conditioning? The study in (Grewe et al., 2017) questions the validity of the Hebbian model in establishing associative learning during fear conditioning. There are several critiques we discuss here. The first critique is that Hebbian plasticity does not explain the experimental finding showing that both upregulation and downregulation of stimulus-evoked responses are present between coactive neurons. The upregulation is provided by our model, so the issue is the downregulation, which is not addressed by our model. However, our model highlights that coactivity alone does not create potentiation; the fine timing of the pre- and postsynaptic spikes determines whether there is potentiation or depression. Here, we find that PING networks are instrumental in setting up the fine timing for potentiation. We suggest that networks not connected to produce the PING may undergo depression when coactive.

The second critique raised by (Grewe et al., 2017) is that Hebbian plasticity alone does not explain why most of the cells exhibiting enhanced responses to the CS did not react to the US before fear conditioning. They suggest that neuromodulators may provide a third condition (besides the activity of the pre- and postsynaptic neurons) that changes the plasticity rule. Our model also does not explicitly address this experimental finding since it requires F to be initially activated by US in order for the fear association to be established. We agree that the fear cells described in (Grewe et al. 2017) may be depolarized by the US without reaching the spiking threshold; however, with neuromodulation provided during the fear training, the same input can lead to spiking, enabling the conditions for Hebbian plasticity. Our discussions above about how neuromodulators affect excitability are relevant to this point. We do not exclude that other forms of plasticity may play a role during fear conditioning in cells not initially activated by the US, but this is not the topic of our modeling study.

The third critique raised by (Grewe et al., 2017) is that Hebbian plasticity cannot explain why the majority of cells that were US- and CS-responsive before training have a reduced CS-evoked response afterward. The reduced response happens over multiple exposures of CS without US; this can involve processes similar to those present in fear extinction, which require plasticity in further networks, especially involving the infralimbic cortex (Milad and Quirk, 2002; Burgos-Robles et al., 2007). An extension of our model could investigate such mechanisms. In the fourth critique, (Grewe et al., 2017) suggests that the Hebbian plasticity rule cannot easily account for the reduction of the responses of many CS+-responsive cells, but not of the CS−-responsive cells. We suggest that the circuits involving paradigms similar to fear extinction do not involve the CS- cells.

Overall, we agree with (Grewe et al., 2017) that neuromodulators play a crucial role in fear conditioning, especially in prolonging the US- and CS-encoding activity as discussed in (see section “Assumptions and predictions of the model” in the Discussion), or even participating in changing the details of the plasticity rule. A possible follow-up of our work involves investigating how fear ensembles form and modify through fear conditioning and later stages. This follow-up work may involve using a tri-conditional rule, as suggested in (Grewe et al., 2017), in which the potential role of neuromodulators is taken into account in the plasticity rule in addition to the pre- and postsynaptic neuron activity. Another direction is to investigate a possible relationship between neuromodulation and a depression-dominated Hebbian rule.”

Finally, we made additional minor changes to the manuscript:

(1) In the Result section “Interneurons interact to modulate fear neuron output”, we specified the following:

“The US input on the pyramidal cell and VIP interneuron is modeled as a Poisson spike train at ~ 50 Hz and an applied current, respectively. In the rest of the paper, we will use the words “US” as shorthand for “the effects of US”.”

(2) In the Result section “Interneuron rhythms provide the fine timing needed for depression dominated STDP to make the association between CS and fear”, we also reported the following:

“Similarly to the US, in the rest of the paper, we will use the words “CS” as shorthand for “the effects of CS”. In our simulations, CS is modeled as a Poisson spike train at ~ 50 Hz, independent of the US input. Thus, we hypothesize that the time structure of the inputs sometimes used for the training (e.g., a series of auditory pips) is not central to the formation of the plasticity in the network.”

Reviewer #2 (Public Reviews):

The authors of this study have investigated how oscillations may promote fear learning using a network model. They distinguished three types of rhythmic activities and implemented an STDP rule to the network aiming to understand the mechanisms underlying fear learning in the BLA.

After the revision, the fundamental question, namely, whether the BLA networks can or cannot intrinsically generate any theta rhythms, is still unanswered. The author added this sentence to the revised version: "A recent experimental paper, (Antonoudiou et al., 2022), suggests that the BLA can intrinsically generate theta oscillations (3-12 Hz) detectable by LFP recordings under certain conditions, such as reduced inhibitory tone." In the cited paper, the authors studied gamma oscillations, and when they applied 10 uM Gabazine to the BLA slices observed rhythmic oscillations at theta frequencies. 10 uM Gabazine does not reduce the GABA-A receptor-mediated inhibition but eliminates it, resulting in rhythmic populations burst driven solely by excitatory cells. Thus, the results by Antonoudiou et al., 2022 contrast with, and do not support, the present study, which claims that rhythmic oscillations in the BLA depend on the function of interneurons. Thus, there is still no convincing evidence that BLA circuits can intrinsically generate theta oscillations in intact brain or acute slices. If one extrapolates from the hippocampal studies, then this is not surprising, as the hippocampal theta depends on extrahippocampal inputs, including, but not limited to the entorhinal afferents and medial septal projections (see Buzsaki, 2002). Similarly, respiratory related 4 Hz oscillations are also driven by extrinsic inputs. Therefore, at present, it is unclear which kind of physiologically relevant theta rhythm in the BLA networks has been modelled.

In our public reply to the Reviewer’s point, we reported the following:

(1) We kindly disagree that (Antonoudiou et al., 2022) contrasts with our study. (Antonoudiou et al., 2022) is a slice study showing that the BLA theta power (3-12 Hz) increases with gabazine compared to baseline. With all GABAergic currents omitted due to gabazine, the LFP is composed of excitatory currents and intrinsic currents. In our model, the high theta (6-12 Hz) comes from the spiking activity of the SOM cells, which increase their activity if the inhibition from VIP cells is removed. Thus, the model produces high theta in the presence of gabazine (see Fig. 1 in our replies to the Reviewers’ public comments). The model also shows that a PING rhythm is produced without gabazine, and that this rhythm goes away with gabazine because PING requires feedback inhibition from PV to fear cells. Thus, the high theta increase and gamma reduction with gabazine in the (Antonoudiou et al., 2022) paper can be reproduced in our model.

(2) We agree that (Antonoudiou et al., 2022) alone is not sufficient evidence that the BLA can produce low theta (3-6 Hz); we discussed a new paper (Bratsch-Prince et al., 2024) that provides further evidence of BLA ability to produce low theta and under what circumstances. The authors reported that intrinsic BLA theta is produced in slices with ACh stimulation (without needing external glutamate input) which, in vivo, would be provided by the basal forebrain (Rajebhosale et al., eLife, 2024) in response to salient stimuli. The low theta depends on muscarinic activation of CCK interneurons, a group of interneurons that overlaps with the VIP neurons in our model (Krabbe 2017; Mascagni and McDonald, 2003). We suspect that the low theta produced in (Bratsch-Prince et al., 2024) is the same as the low theta in our model. In future work, we will aim to show that ACh activates the BLA VIP cells, which are essential to the low theta generation in the network.

In the manuscript, we added to and modified the Discussion section “Where the rhythms originate, and by what mechanisms”. This text aims to better discuss (Antonoudiou et al. 2022) and introduce (Bratsch-Prince et al., 2024) with its connection to our hypothesis that the theta oscillations can be produced within the BLA. The new version is:

“Where the rhythms originate, and by what mechanisms. A recent experimental paper (Antonoudiou et al., 2022) suggests that the BLA can intrinsically generate theta oscillations (312 Hz) detectable by LFP recordings when inhibition is totally removed due to gabazine application. They draw this conclusion in mice by removing the hippocampus, which can volume conduct to BLA, and noticing that other nearby brain structures did not display any oscillatory activity. In our model, we note that when inhibition is removed, both AMPA and intrinsic currents contribute to the network dynamics and the LFP. Thus, interneurons with their specific intrinsic currents (i.e., D-current in the VIP interneurons, and NaP- and H- currents in SOM interneurons) can indeed affect the model LFP and support the generation of theta and gamma rhythms (Fig. 6G).

Another slice study, (Bratsch-Prince et al., 2024), shows that BLA is intrinsically capable of producing a low theta rhythm with ACh stimulation and without needing external glutamate input. ACh is produced in vivo by the basal forebrain in response to US (Rajebhosale et al., 2024). Although we did not explicitly include the BF and ACh modulation of BLA in our model, we implicitly include the effect of ACh in BLA by increasing the activity of the VIP cells, which then produce the low theta rhythm. Indeed, low theta in the BLA is known to depend on the muscarinic activation of CCK interneurons, a group of interneurons that overlaps with the class of VIP neurons in our model (Mascagni and McDonald, 2003; Krabbe et al., 2018).

Although the BLA can produce these rhythms, this does not rule out that other brain structures also produce the same rhythms through different mechanisms, and these can be transmitted to the BLA. Specifically, it is known that the olfactory bulb produces and transmits the respiratoryrelated low theta (4 Hz) oscillations to the dorsomedial prefrontal cortex, where it organizes neural activity (Bagur et al., 2021). Thus, the respiratory-related low theta may be captured by BLA LFP because of volume conduction or through BLA extensive communications with the prefrontal cortex. Furthermore, high theta oscillations are known to be produced by the hippocampus during various brain functions and behavioral states, including during spatial exploration (Vanderwolf, 1969) and memory formation/retrieval (Raghavachari et al., 2001), which are both involved in fear conditioning. Similarly to the low theta rhythm, the hippocampal high theta can manifest in the BLA. It remains to understand how these other rhythms may interact with the ones described in our paper. However, we emphasize that there is also evidence (as discussed above) that these rhythms arise within the BLA.”

Reviewer #2 (Recommendations for the Authors):

(1) Three different types of VIP interneurons with distinct firing patterns have been revealed in the BLA (Rhomberg et al., 2018). Does the generation of rhythmic activities depend on the firing features of VIP interneurons? Does it matter whether VIP interneurons fire burst of action potentials or they discharge more regularly?

(2) The authors used data for modeling SST interneurons obtained e.g., in the hippocampus. However, there are studies in the BLA where the intrinsic characteristics of SST interneurons have been reported (Unal et al., 2020; Guthman et al., 2020; Vereczki et al., 2021). Have the authors considered using results of studies that were conducted in the BLA?

We thank the Reviewer for their questions, which have helped us further improve our manuscript in response to similar queries from Reviewer 3 in the previous review round. More in detail:

(1) Although other electrophysiological types exist (Sosulina et al., 2010), we hypothesized that the electrophysiological type of VIP neurons that display intrinsic stuttering is the type that would be involved in mediating low theta oscillations during fear conditioning. This is because VIP intrinsic stuttering in cortical neurons is thought to involve the D-current, which helps create low theta bursting oscillations in the neuronal spiking patterns (Chartove et al., 2020). We think that the other subtypes of VIP interneurons are not essential for the low theta oscillatory dynamics observed during fear conditioning and, thus, did not provide an essential constraint for the phenomena we are trying to capture. VIP interneurons in our network must fire bursts at low theta to be effective in creating the pauses in ECS and F spiking needed for potentiation; single spikes at theta are not sufficient to create these pauses.

(2) In our model, we used the results conducted in a BLA study (Sosulina et al., 2010). SOM cells in the BLA display several physiologic types. We chose to include in our model the type showing early adaptation in response to a depolarizing current and inward (outward) rectification upon the initiation (release) of a hyperpolarizing current. We hypothesize that this type can produce high theta oscillations, a prominently observed rhythm in the BLA. Unal et al., 2020 (Unal et al., 2020) found two populations of SOM cells in the BLA, which have been previously recorded in (Sosulina et al., 2010), including the one type we chose to model. This SOM cell type shows a low threshold spiking profile characterized by spike frequency adaptation and voltage sag indicative of an H-current used in our model. Guthman et al., 2020, (Guthman et al., 2020), also found a population of SOM cells with hyperpolarization induced sag.

Our model also uses a NaP-current for which there is no data in the BLA. However, it is known to exist in hippocampal SOM cells and that NaP- and H- currents can produce such a high theta in hippocampal cells. It is a standard practice in modeling to use the best possible replacement for unknown currents. Of course, it is unfortunate to have to do this. We also note that models can be considered proof of principle, that can be proved or disproved by further experimental work. Both (Guthman et al., 2020) and (Vereczki et al., 2021) also uncover further heterogeneity among BLA SOM interneurons involving more than electrophysiology. We hypothesize that such a level of heterogeneity revealed by these three studies is not key to the question we are asking (where crucial ingredients are the rhythms) and, therefore, was not included in our minimal model.

We modified the Discussion section titled “Assumptions and predictions of the model” as follows:

“Our model, which is a first effort towards a biophysically detailed description of the BLA rhythms and their functions, does not include the neuron morphology, many other cell types, conductances, and connections that are known to exist in the BLA; models such as ours are often called “minimal models” and constitute most biologically detailed models. For example, although there is considerable variability in the activity patterns of both VIP cells and SOM cells (Sosulina et al., 2010; Guthman et al., 2020; Ünal et al., 2020; Vereczki et al., 2021), our focus was specifically on those subtypes that generate critical rhythms within the BLA. Such minimal models are used to maximize the insight that can be gained by omitting details whose influence on the answers to the questions addressed in the model are believed not to be qualitatively important. We note that the absence of these omitted features constitutes hypotheses of the model: we hypothesize that the absence of these features does not materially affect the conclusions of the model about the questions we are investigating. Of course, such hypotheses can be refuted by further work showing the importance of some omitted features for these questions and may be critical for other questions. Our results hold when there is some degree of heterogeneity of cells of the same type, showing that homogeneity is not a necessary condition.”

(3) The authors may double-check the reference list, as e.g., Cuhna-Reis et al., 2020 is not listed.

We thank the Reviewer for spotting this. We checked the reference list and all the references are now listed.

Finally, we wanted to acknowledge that we made other changes to the manuscript unrelated to the reviewers’ questions with the purpose of gaining clarity. More specifically:

(1) We included a section titled “Significance” after the abstract and keywords, which reads as follows:

“Our paper accounts for the experimental evidence showing that amygdalar rhythms exist, suggests network origins for these rhythms, and points to their central role in the mechanisms of plasticity involved in associative learning. It is one of the few papers to address high-order cognition with biophysically detailed models, which are sometimes thought to be too detailed to be adequately constrained. Our paper provides a template for how to use information about brain rhythms to constrain biophysical models. It shows in detail, for the first time, how multiple interneurons help to provide time scales necessary for some kinds of spike-timing-dependent plasticity (STDP). It spells out the conditions under which such interactions between interneurons are needed for STDP and why. Finally, our work helps to provide a framework by which some of the discrepancies in the fear learning literature might be reevaluated. In particular, we discuss issues about Hebbian plasticity in fear learning; we show in the context of our model how neuromodulation might resolve some of those issues. The model addresses issues more general than that of fear learning since it is based on interactions of interneurons that are prominent in the cortex, as well as the amygdala.”

(2) The Result section “Physiology of the interneuron types is critical to their role in depression-dominated plasticity”, which is now titled “Mechanisms by which interneurons contribute to potentiation in depression-dominated plasticity”, now reads as follows:

“Mechanisms by which interneurons contribute to potentiation during depressiondominated plasticity. The PV cell is necessary to induce the correct pre-post timing between ECS and F needed for long-term potentiation of the ECS to F conductance. In our model, PV has reciprocal connections with F and provides lateral inhibition to ECS. Since the lateral inhibition is weaker than the feedback inhibition, PV tends to bias ECS to fire before F. This creates the fine timing needed for the depression-dominated rule to instantiate plasticity. If we used the classical Hebbian plasticity rule (Bi and Poo, 2001) with gamma frequency inputs, this fine timing would not be needed and ECS to F would potentiate over most of the gamma cycle, and thus we would expect random timing between ECS and F to lead to potentiation (Fig. S4). In this case, no interneurons are needed (See Discussion “Synaptic plasticity in our model” for the potential necessity of the depression-dominated rule).

In this network configuration, the pre-post timing for ECS and F is repeated robustly over time due to coordinated gamma oscillations (PING, as shown in Fig. 4A, Fig. 1C) arising through the reciprocal interactions between F and PV (Feng et al., 2019). PING can arise only when PV is in a sufficiently low excitation regime such that F can control PV activity (Börgers et al., 2005), as in Fig. 4A. However, although such a low excitation regime establishes the correct fine timing for potentiation, it is not sufficient to lead to potentiation (Fig. 4A, Fig. S2C): the depression-dominated rule leads to depression rather than potentiation unless the PING is periodically interrupted. During the pauses, made possible only in the full network by the presence of VIP and SOM, the history-dependent build-up of depression decays back to baseline, allowing potentiation to occur on the next ECS/F active phase. (The detailed mechanism of how this happens is in the Supplementary Information, including Fig. S2). Thus, a network without the other interneuron types cannot lead to potentiation. Though a low excitation level for a PV cell is necessary to produce a PING, a higher excitation level is necessary to produce a pause in the ECS and F. This higher excitation level is consistent with the experimental literature showing a strong activation of PV after the onset of CS (Wolff et al., 2014). The higher excitation happens when the VIP cell is silent, whereas a low excitation level is achieved when the VIP cell fires and partially inhibits the PV cell (Fig. 4B, Fig. S2D). The interruption in the ECS and F activity requires the participation of another interneuron, the SOM cell (Figs. 2B, S2): the pauses in inhibition from the VIP periodically interrupt ECS and F firing by releasing PV and SOM from inhibition and thus indirectly silencing ECS and F. Without these pauses, depression dominates (see SI section “ECS and F activity patterns determine overall potentiation or depression”).”

We also removed a supplementary figure (Fig. S2).

(3) We wanted to be clear and motivate our choice to extend the low theta range to 2-6 Hz and the high theta range to 6-14 Hz, compared to the 3-6 Hz and 6-12 Hz, respectively in the BLA experimental literature. Our main reason for extending the ranges was because the peaks of low and high theta power in the VIP and SOM cells, respectively, (the cells that generate these oscillations) occurred at the borders of the experimental ranges. Thus, in order to include the peaks of the model LFP, we lowered the low theta range by 1 Hz and increased the high theta range by 2 Hz.

We present a new supplementary figure (Fig. S1) containing the power spectra of VIP, which is the source of low theta in our model, and SOM interneuron, which is the source of high theta:

We mention Fig. S1 in the Result section “Rhythms in the BLA can be produced by interneurons”, where we added the following text: o “In the baseline condition, the condition without any external input from the fear conditioning paradigm (Fig. 1B, top), our VIP neurons exhibit short bursts of gamma activity (~38 Hz) at low theta frequencies (~2-6 Hz) (peaking at ~3.5 Hz) (see Fig. S1A).” o “In our baseline model, SOM cells have a natural frequency of ~12 Hz (Fig. 1B, middle; Fig. S1B), which is at the upper limit of the experimental high theta range; this motivates our choice to extend the high theta range up to 14 Hz in order to include the peak.”

Knowing the natural frequencies of VIP and SOM interneurons from the Result section “Rhythms in the BLA can be produced by interneurons”, we specified more clearly that we quantify the change of power in the low and high theta range around the power peaks in those ranges. Specifically, we changed some sentences in the first paragraph of the Result section “Increased low-theta frequency is a biomarker of fear learning” as follows:

“We find that fear conditioning leads to an increase in low theta frequency power of the network spiking activity compared to the pre-conditioned level (Fig. 6 A,B); there is no change in the high theta power. We also find that the LFP, modeled as the linear sum of all the AMPA, GABA, NaP-, D-, and H- currents in the network, similarly reveals a low theta power increase when considering the peak of the low theta power, and no significant variation in the high theta power again when considering the peak of the high theta power (Fig. 6 C,D,E).”

Finally, we made a few other small changes:

In the Introduction, we mention the following: “We also note that there is not uniformity on the exact frequencies associated with low and high theta, e.g., ((Lorétan et al., 2004) used 2-6 Hz for low theta). Here, we use 2-6 Hz for the theta range and 6-14 Hz for the high theta range.”

In Fig. 6DE (reported below point 3)), we reran the statistics using a smaller interval for high theta (11.5-13 Hz) to focus around the peak. Our initial result showing significant change in low theta between pre and post fear conditioning and no change in high theta still holds.

In Fig. 6 of the Result section “Increase low-theta frequency is a biomarker of fear learning”, we switched the order of panels F and G. This change allows us to first focus on the AMPA currents, which are the major contributors of the low theta power increase, and to specify what AMPA current drives that increase. After that, we present the power spectrum of the GABA currents, as well.

The corresponding text in the Result section, now reads as follows:

“We find that fear conditioning leads to an increase in low theta frequency power of the network spiking activity compared to the pre-conditioned level (Fig. 6 A,B); there is no change in the high theta power. We also find that the LFP, modeled as the linear sum of all the AMPA, GABA, NaP-, D-, and H- currents in the network, similarly reveals a low theta power increase when considering the peak of the low theta power, and no significant variation in the high theta power again when considering the peak of the high theta power (Fig. 6 C,D,E). These results are consistent with the experimental findings in (Davis et al., 2017). Specifically, the newly potentiated AMPA synapse from ECS to F ensures F is active after fear conditioning, thus generating strong currents in the PV cells to which it has strong connections (Fig. 6F). It is the AMPA currents to the PV interneurons that are directly responsible for the low theta increase; it is the newly potentiated ECS to F synapse that paces the AMPA currents in the PV interneurons to go at low theta. Thus, the low theta increase is due to added excitation provided by the new learned pathway.”

(4) In the Discussion section “Assumptions and predictions of the model”, we specified the following:

“Our model predicts that blockade of D-current in VIP interneurons (or silencing VIP interneurons) will both diminish low theta and prevent fear learning. Finally, the model assumes the absence of significantly strong connections from the excitatory projection cells ECS to PV interneurons, unlike the ones from F to PV. Including those synapses would alter the PING rhythm created by the interactions between F and PV, which is crucial for fine timing between ECS and F needed for LTP.”

(5) Finally, to broaden the potential interest of our study, we added the following sentences:

At the conclusion of the abstract:

“The model makes use of interneurons commonly found in the cortex and, hence, may apply to a wide variety of associative learning situations.” - At the conclusion of the introduction:

“Finally, we note that the ideas in the model may apply very generally to associative learning in the cortex, which contains similar subcircuits of pyramidal cells and interneurons: PV, SOM and VIP cells.”

Also, changes in the emphasis of the paper led us to remove the following from the abstract: “Finally, we discuss how the peptide released by the VIP cell may alter the dynamics of plasticity to support the necessary fine timing.”

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The basolateral amygdala (BLA) is a key site where fear learning takes place through synaptic plasticity. Rodent research shows prominent low theta (∼3-6 Hz), high theta (∼6-12 Hz), and gamma (>30 Hz) rhythms in the BLA local field potential recordings. However, it is not understood what role these rhythms play in supporting the plasticity. Here, we create a biophysically detailed model of the BLA circuit to show that several classes of interneurons (PV, SOM, and VIP) in the BLA can be critically involved in producing the rhythms; these rhythms promote the formation of a dedicated fear circuit shaped through rhythmic gating of spike-timing-dependent plasticity. Each class of interneurons is necessary for the plasticity. We find that the low theta rhythm is a biomarker of successful fear conditioning. Finally, we discuss how the peptide released by the VIP cell may alter the dynamics of plasticity to support the necessary fine timing.

Revised version: The basolateral amygdala (BLA) is a key site where fear learning takes place through synaptic plasticity. Rodent research shows prominent low theta (∼3-6 Hz), high theta (∼6-12 Hz), and gamma (>30 Hz) rhythms in the BLA local field potential recordings. However, it is not understood what role these rhythms play in supporting the plasticity. Here, we create a biophysically detailed model of the BLA circuit to show that several classes of interneurons (PV, SOM, and VIP) in the BLA can be critically involved in producing the rhythms; these rhythms promote the formation of a dedicated fear circuit shaped through spike-timing-dependent plasticity. Each class of interneurons is necessary for the plasticity. We find that the low theta rhythm is a biomarker of successful fear conditioning. The model makes use of interneurons commonly found in the cortex and, hence, may apply to a wide variety of associative learning situations.

Introduction

Earlier version: Pavlovian fear conditioning is widely used as a model of associative learning across multiple species (Phelps and LeDoux, 2005), and has been used more generally to study plasticity and memory formation. The major questions, still under investigation, relate to how plasticity is instantiated in defined circuits and how it is regulated by the circuit components (Rumpel et al., 2005; Sah et al., 2008; Johansen et al., 2014; Bocchio et al., 2017; Grewe et al., 2017). The fear conditioning paradigm consists of a neutral stimulus (conditioned stimulus, CS) presented one or more times together with an aversive stimulus (unconditioned stimulus, US), which induces a fear response.

Revised version: Pavlovian fear conditioning is widely used as a model of associative learning across multiple species (Phelps and LeDoux, 2005) and has been used more generally to study plasticity and memory formation. The major open questions relate to how plasticity is instantiated in defined circuits and how it is regulated by the circuit components (Rumpel et al., 2005; Sah et al., 2008; Johansen et al., 2014; Bocchio et al., 2017; Grewe et al., 2017). The fear conditioning paradigm consists of a neutral stimulus (conditioned stimulus, CS) presented one or more times together with an aversive stimulus (unconditioned stimulus, US), which induces a fear response. The presentation of the US and CS may or may not overlap in time (e.g., see (Laxmi et al., 2003; Stujenske et al., 2014; Krabbe et al., 2019)).

Earlier version: In the basolateral amygdala (BLA), the main site of fear learning in the mammalian brain (Fanselow and LeDoux, 1999; Tovote et al., 2015; Krabbe et al., 2018), local field potential recordings (LFP) show prominent low theta (∼3-6 Hz), high theta (∼6-12 Hz), and gamma (>30 Hz) rhythms (Seidenbecher et al., 2003; Courtin et al., 2014b; Stujenske et al., 2014; Davis et al., 2017). A recent rodent experimental study (Antonoudiou et al., 2022) suggests that BLA can intrinsically generate theta oscillations (3-12 Hz). Furthermore, other studies show increased low theta (Davis et al., 2017) and gamma (Courtin et al., 2014b) in local field potential (LFP) recordings after successful fear conditioning, whereas modulation of high theta is associated with fear extinction (Davis et al., 2017), a paradigm that aims to suppress the association between CS and fear (Bouton, 2004); the modulation of the power of these rhythms suggests they may be associated with BLA plasticity.

Revised version: In the basolateral amygdala (BLA), the main site of fear learning in the mammalian brain (Fanselow and LeDoux, 1999; Tovote et al., 2015; Krabbe et al., 2018), local field potential recordings (LFP) show prominent low theta (∼3-6 Hz), high theta (∼6-12 Hz), and gamma (>30 Hz) rhythms (Seidenbecher et al., 2003; Courtin et al., 2014b; Stujenske et al., 2014; Davis et al., 2017). Recent rodent experimental studies (Antonoudiou et al., 2022; Bratsch-Prince et al., 2024) suggest that BLA can intrinsically generate theta oscillations (3-12 Hz). Furthermore, other studies show increased low theta (Davis et al., 2017) and gamma (Courtin et al., 2014b) in LFP recordings after successful fear conditioning, whereas modulation of high theta is associated with fear extinction (Davis et al., 2017), a paradigm that aims to suppress the association between CS and fear (Bouton, 2004); the modulation of the power of these rhythms suggests they may be associated with BLA plasticity. However, the potential roles of rhythms in instantiating plasticity needed for successful learning are still under investigation (Bocchio et al., 2017). We also note that there is not uniformity on the exact frequencies associated with low and high theta, e.g., ((Lorétan et al., 2004) used 2-6 Hz for low theta). Here, we use 2-6 Hz for the theta range and 6-14 Hz for the high theta range.

Earlier version: The potential roles of rhythms in instantiating the plasticity needed for successful learning is still under investigation (Bocchio et al., 2017). Fear responses are mediated by a network of several brain areas including, but not limited to, the BLA, the prefrontal cortex (mPFC), and the hippocampus (Seidenbecher et al., 2003; Stujenske et al., 2014; Tovote et al., 2015). In experimental studies, successful fear conditioning leads fear-encoding neurons to respond to the CS. In our work, it results in strengthening the connection between neurons encoding CS and neurons encoding fear. We create a biophysically detailed model of the BLA circuit suggesting that interneurons in the BLA can be critically involved in producing the experimentally measured rhythms and are mechanistically important in the instantiation of plasticity during learning. Indeed, we find that the low theta, fast theta, and gamma rhythms in the BLA originating from the BLA interneurons promote the formation of a dedicated fear circuit shaped through rhythmic regulation of depression-dominated spike-timing-dependent plasticity (STDP). (See Discussion about non-depression-dominated plasticity related to fear conditioning.) In this model, if any of the classes of interneurons are removed from the circuit, the rhythms are changed and there is a failure of the plasticity needed for successful learning. We show that fine timing between excitatory projection neurons responding to the CS and US is necessary but not sufficient to produce associative plasticity when the latter can be affected by all the spikes, not just nearest neighbors. The other critical element for plasticity is the interaction of interneurons that creates pauses in excitatory cell activity.

Revised version: In this paper, we aim to show 1) How a variety of BLA interneurons (PV, SOM and VIP) lead to the creation of these rhythms and 2) How the interaction of the interneurons and the rhythms leads to the appropriate timing of the cells responding to the US and those responding to the CS to promote fear association through spike-timing-dependent plasticity (STDP). Since STDP requires overlap of the effects of the CS and US, and some conditioning paradigms do not have overlapping US and CS, we include as a hypothesis that the effects of the CS and US overlap even if the CS and US stimuli do not. In the Discussion, we suggest how neuromodulation by ACh and/or dopamine can provide such overlap. We create a biophysically detailed model of the BLA circuit involving all three types of interneurons and show how each may participate in producing the experimentally observed rhythms and interacting to produce the necessary timing for fear learning. In particular, we find that the low theta, high theta, and gamma rhythms in the BLA originating from the BLA interneurons promote the formation of a dedicated fear circuit shaped through rhythmic regulation of depression-dominated STDP. In this model, if any of the classes of interneurons are removed from the circuit, the rhythms are changed and there is a failure of the plasticity needed for successful learning. We show that fine timing between excitatory projection neurons responding to the CS and US is necessary but not sufficient to produce associative plasticity when the latter can be affected by the entire history of spikes, not just the most recent pre- and post-synaptic spikes. The other critical element for plasticity is the interaction of interneurons that creates pauses in excitatory cell activity.

Earlier version: The model reproduces the increase in the low theta after training that was found in the experimental data. This increase was not seen for network instantiations that did not learn; thus, the simulations suggest that the increase in low theta is a biomarker of successful fear conditioning. Furthermore, this low theta signal emanates from the newly formed synapse linking CS- and fear-encoding neurons. These results are also highly relevant to the recent study (Veit et al., 2023) in which the same types of interneurons used in this paper are shown to allow global selectivity in coherence in gamma oscillations (see Discussion).

Revised version: The model reproduces the increase in the low theta after training that was found in the experimental data (Davis et al., 2017). This increase was not seen for network instantiations that did not learn; thus, the simulations suggest that the increase in low theta is a biomarker of successful fear conditioning. Furthermore, this low theta signal is a consequence of network dynamics that emerge due to the newly formed synapse linking CS-encoding and fear-encoding neurons. Finally, we note that the ideas in the model may apply very generally to associative learning in the cortex, which contains similar subcircuits of pyramidal cells and interneurons: PV, SOM and VIP cells.

Results

Rhythms in the BLA can be produced by interneurons

Brain rhythms are thought to be generated and propagated largely by interneurons (Whittington et al., 2000). Identified interneurons in BLA include VIP (vasoactive intestinal peptide-expressing), SOM (somatostatin-expressing), and PV (parvalbumin-expressing) (Muller et al., 2006, 2007; Rainnie et al., 2006; Bienvenu et al., 2012; Krabbe et al., 2018, 2019), which can be further subdivided according to their electrophysiological dynamics (Sosulina et al., 2010; Spampanato et al., 2011). In the model, we show that some types of VIP, SOM, and PV can each contribute to the generation of a key rhythm involved in the BLA due to their specific intrinsic currents.

Earlier version: For VIP interneurons, we consider a subtype that responds to a depolarizing step current with bursting or stuttering behavior (Sosulina et al., 2010; Spampanato et al., 2011). This type of behavior can be elicited by a D-type potassium current, a current thought to be found in a similar electrophysiological subtype of VIP interneurons in the cortex (Porter et al., 1998). In our model, VIP interneurons endowed with a D-current respond to depolarizing currents with long bursts, and to hyperpolarizing currents with no action potentials, thus reproducing the electrophysiological properties of type I BLA interneurons in (Sosulina et al., 2010) (Fig. 1A, top): our VIP neurons exhibit gamma (∼52 Hz) bursting activity at low theta frequencies (∼3-6 Hz) in the baseline condition, the condition without any external input from the fear conditioning paradigm (Fig. 1B, top).

Revised version: For VIP interneurons, we consider a subtype that responds to a depolarizing step current with bursting or stuttering behavior (Sosulina et al., 2010; Spampanato et al., 2011). This type of behavior can be elicited by a D-type potassium current, a current thought to be found in a similar electrophysiological subtype of VIP interneurons in the cortex (Porter et al., 1998). In our model, VIP interneurons endowed with a D-current respond to depolarizing currents with long bursts, and to hyperpolarizing currents with no action potentials, thus reproducing the electrophysiological properties of type I BLA VIP interneurons in (Sosulina et al., 2010) (Fig. 1A, top): in the baseline condition, the condition without any external input from the fear conditioning paradigm (Fig. 1B, top), our VIP neurons exhibit short bursts of gamma activity (∼38 Hz) at low theta frequencies (∼2-6 Hz) (peaking at ∼3.5 Hz) (see Fig. S1A).

Isolated neurons produce fundamental rhythms.

A. dynamics in response to depolarizing and hyperpolarizing currents mimic the electrophysiological behavior of BLA interneurons classified in (Sosulina et al., 2010). B. dynamics in the baseline condition. C. Interacting PV and excitatory projection neuron (E) entrain in a pyramidal-interneuron network gamma rhythm (PING).

Earlier version: SOM interneurons have been well studied in the hippocampus, especially the O-LM cells (Maccaferri and McBain, 1996; Gillies et al., 2002; Saraga et al., 2003; Rotstein et al., 2005). Although the intrinsic currents of BLA SOM cells are unknown, SOM cells in the hippocampus have a hyperpolarization-activated current, i.e., H-current, and a persistent sodium current, i.e., NaP-current. In our model, with the introduction of these currents with specific conductances (see Materials and Methods for details), the SOM cells mimic the electrophysiologic behavior of type III BLA SOM cells in (Sosulina et al., 2010), showing regular spikes with early spike-frequency adaptation in response to a depolarizing current, and pronounced inward rectification (downward deflection) and outward rectification (upward deflection) upon the initiation and release of a hyperpolarizing current (Fig. 1A, middle). In our baseline model, SOM cells have a natural frequency rhythm in the high theta range (6-12 Hz) (Fig. 1B, middle).

Revised version: For SOM interneurons, we focus on the electrophysiologic behavior of type III BLA SOM cells in (Sosulina et al., 2010), showing regular spikes with early spike-frequency adaptation in response to a depolarizing current, and pronounced inward rectification (downward deflection) and outward rectification (upward deflection) upon the initiation and release of a hyperpolarizing current. SOM interneurons have been mostly studied in the hippocampus, especially the O-LM cells (Maccaferri and McBain, 1996; Gillies et al., 2002; Saraga et al., 2003; Rotstein et al., 2005), which are known to have a hyperpolarization-activated current, i.e., H-current, and a persistent sodium current, i.e., NaP-current. An H-current has also been observed in the SOM cells in the BLA (Ünal et al., 2020). In our model, with the introduction of both H- and NaP-currents with specific conductances (see Materials and Methods for details), the SOM cells mimic the electrophysiologic behavior of type III BLA SOM cells in (Sosulina et al., 2010) (Fig. 1A, middle). In our baseline model, SOM cells have a natural frequency of ∼12 Hz (Fig. 1B, middle; Fig. S1B), which is at the upper limit of the experimental high theta range; this motivates our choice to extend the high theta range up to 14 Hz in order to include the peak.

Earlier version: Our model PV interneurons are fast-spiking interneurons (FSIs) with standard action potentials produced by Hodgkin-Huxley-type sodium and potassium conductances. They are silent at baseline condition and show similar behaviors to type IV interneurons (Sosulina et al., 2010) in response to depolarizing and hyperpolarizing currents (Fig. 1A, bottom). At baseline, our PV interneurons are silent (Fig. 1B, bottom). However, when reciprocal connections are present and the excitatory projection neuron receives enough excitation to fire, the PV interneuron forms a PING rhythm (pyramidal-interneuron network gamma) (Whittington et al., 2000) (Fig. 1C); this has been suggested as a possible mechanism for the basis of gamma rhythm generation in the BLA (Feng et al., 2019). The reciprocal interactions between the projection neuron and PV form a gamma rhythm that depends sensitively on the external input to the excitatory projection neuron and the PV’s decay time constant of inhibition (8.3 ms).

Revised version: Our model PV interneurons are fast-spiking interneurons (FSIs) with standard action potentials produced by Hodgkin-Huxley-type sodium and potassium conductances. They are silent at baseline condition (Fig. 1B, bottom) and show similar behaviors to type IV interneurons (Sosulina et al., 2010) in response to depolarizing and hyperpolarizing currents (Fig. 1A, bottom). However, when reciprocal connections are present between the PV interneuron and the excitatory projection neuron, these neurons form a PING rhythm (pyramidal-interneuron network gamma) (Whittington et al., 2000), if the excitatory projection neuron receives enough excitation to fire (Fig. 1C); this has been suggested as a possible mechanism for the basis of gamma rhythm generation in the BLA (Feng et al., 2019). The frequency of PING depends sensitively on the external input to the excitatory projection neuron and the PV’s decay time constant of inhibition (8.3 ms).

Interneurons interact to modulate fear neuron output

Earlier version: Our BLA network consists of interneurons, detailed in the previous section, and excitatory projection neurons (Fig. 2A). Both the fear-encoding neuron (F), an excitatory projection neuron, and the VIP interneuron are activated by the noxious stimulus US (Krabbe et al., 2019). As shown in Fig. 2A (top, right), VIP disinhibits F by inhibiting both SOM and PV, as suggested in (Krabbe et al., 2019). We do not include connections from PV to SOM and VIP, nor connections from SOM to PV and VIP, since those connections have been shown to be significantly weaker than the ones included (Krabbe et al., 2019). The simplest network we consider is made of one neuron for each cell type. We introduce some heterogeneity in the last two sections of the Results.

Revised version: Our BLA network consists of interneurons, detailed in the previous section, and excitatory projection neurons (Fig. 2A). Both the fear-encoding neuron (F), an excitatory projection neuron, and the VIP interneuron are activated by the noxious stimulus US (Krabbe et al., 2019). The US input on the pyramidal cell and VIP interneuron is modeled as a Poisson spike train at ∼ 50 Hz and an applied current, respectively. In the rest of the paper, we will use the words “US” as shorthand for “the effects of US”. As shown in Fig. 2A (top, right), VIP disinhibits F by inhibiting both SOM and PV, as suggested in (Krabbe et al., 2019). We do not include connections from PV to SOM and VIP, nor connections from SOM to PV and VIP, since those connections have been shown to be significantly weaker than the ones included (Krabbe et al., 2019). The simplest network we consider is made of one neuron for each cell type. We introduce additional neurons for each cell type with some heterogeneity in the last two sections of the Results.

BLA interneurons and the excitatory projection neurons interact and modulate the network activity.

A (top): Network made of three interneurons (VIP, SOM, and PV) and the excitatory projection neuron encoding fear (F) without US input (left) and with US input (right). A (bottom): before the onset of US, VIP shows gamma bursts nested in the low theta rhythm (blue trace), and SOM fires at a natural frequency in the high theta range (purple trace). PV is completely silent due to the lack of any external input (green trace). F, despite its natural frequency of around 11 Hz, is silent due to the inhibition from SOM (orange trace). After US onset, due to the longer VIP bursts and the US input, F shows a pronounced activity during the VIP active phase and outside when the SOM and PV inhibition fade. PV is active only when excited by F, and then gives inhibitory feedback to F. B (top): Network in panel A with the excitatory projection neuron encoding the CS input (ECS) during the CS presentation (left) and with paired CS and US inputs (right). B (bottom): 2-second dynamics of all the neurons in the BLA network affected by CS, and by US after 1 second has elapsed. As in panel A, VIP shows gamma bursting activity nested in the low theta frequency range with bursts duration affected by the presence or absence of US. VIP inhibits i) SOM, which fires at high theta (purple trace) regardless of the external inputs, and ii) PV, which fires at gamma. ECS (light blue trace) and F are both active when both CS and US are present and VIP is active, producing a gamma nested into a low theta rhythm. The evolution in time of the conductance (gAMPA) shows an overall potentiation over the second half of the dynamics when both ECS and F are active. C: blowup of ECS-F burst of activity and gAMPA dynamic shown in the gray area in panel B (bottom); ECS (blue trace) fires most of the time right before F, thus creating the correct pre-post timing conducive for potentiation of the ECS to F conductance. The order of each pair of ECS-F spikes is labeled with “c” (correct) or “w” (wrong).

Earlier version: Fig. 2A (bottom) shows a typical dynamic of the network before and after the US input onset, with US modeled as a Poisson spike train at ∼50 Hz; the network produces all the rhythms originating from the interneurons alone or through their interactions with the excitatory projection neurons (shown in Fig. 1). Specifically, since VIP is active at low theta during both rest and upon the injection of US, it then modulates F at low theta cycles via SOM and PV. In the baseline condition, the VIP interneuron has short gamma bursts nested in low theta rhythm. With US onset, VIP increases its burst duration and the frequency of low theta rhythm. These longer bursts make the SOM cell silent for long periods of each low theta cycle, providing F with windows of disinhibition and contributing to the abrupt increase in activity right after the US onset. Finally, in Fig. 2A, PV lacks any external input and fires only when excited by F. Thanks to their reciprocal interactions, PV forms a PING rhythm with F, as depicted in Fig.1C.

Revised version: Fig. 2A shows a typical dynamic of the network before and after the US input onset; The network produces all the rhythms originating from the interneurons alone or through their interactions with the excitatory projection neuron (shown in Fig. 1). Specifically, since VIP is active at low theta during both rest and upon the injection of US, it then modulates F at low theta cycles via SOM and PV. In the baseline condition, the VIP interneuron has short gamma bursts nested in a low theta rhythm. With US onset, VIP increases its burst duration and the frequency of the low theta rhythm. These longer bursts make the SOM cell silent for long periods of each low theta cycle, providing F with windows of disinhibition and contributing to the abrupt increase in the activity of F right after the US onset. Finally, in Fig. 2A, PV lacks any external input and fires only when excited by F. Due to their reciprocal interactions, PV forms a PING rhythm with F, as depicted in Fig.1C.

Interneuron rhythms provide the fine timing needed for depression-dominated STDP to make the association between CS and fear

Earlier version: We now introduce another excitatory projection neuron (ECS), as shown in Fig. 2B (top). ECS, unlike F, responds to the neutral stimulus CS, as does PV. In our simulations, the CS input describes either the context or the cue in contextual and cued fear conditioning, respectively. For the context, the input may come from the hippocampus or other non-sensory regions, but this does not affect its role as input in the model. By the end of fear conditioning, CS consistently activates the neuron F, thus eliciting the network fear response. This happens because of the formation and strengthening of the synapse from ECS to F by means of synaptic plasticity. We now show how this network, with appropriate connection strengths among neurons, can make the timing of the interneurons confer pre-post timing to ECS and F, which is conducive to spike-timing-dependent plasticity potentiation suggested to be critical for associative aversive learning (Rogan et al., 1997; Nabavi et al., 2014); in particular, we need feedback inhibition (from PV to F) to be stronger than lateral inhibition (from PV to ECS) to promote ECS firing before F. The Hebbian plasticity rule that we use is characterized by a longer time constant of depression than potentiation (and equal maximal amplitudes) and considers the whole history of ECS and F spiking activity (see Materials and Methods and Fig. S1A for more details).

Revised version: We now introduce another excitatory projection neuron (ECS), as shown in Fig. 2B (top). ECS, unlike F, responds to the neutral stimulus CS, as does PV. Similarly to the US, in the rest of the paper, we will use the words “CS” as shorthand for “the effects of CS”. In our simulations, CS is modeled as a Poisson spike train at ∼ 50 Hz, independent of the US input. Thus, we hypothesize that the time structure of the inputs sometimes used for the training (e.g., a series of auditory pips) is not central to the formation of the plasticity in the network. Our CS input describes either the context or the cue in contextual and cued fear conditioning, respectively. For the context, the input may come from the hippocampus or other non-sensory regions, but this does not affect its role as input in the model. By the end of fear conditioning, CS consistently activates the neuron F, thus eliciting the network fear response. This happens because of the formation and strengthening of the synapse from ECS to F by means of synaptic plasticity. We now show how this network, with appropriate connection strengths among neurons, can make the timing of the interneurons confer pre-post timing to ECS and F, which is conducive to spike-timing-dependent plasticity potentiation suggested to be critical for associative aversive learning (Rogan et al., 1997; Nabavi et al., 2014); in particular, we need feedback inhibition (from PV to F) to be stronger than lateral inhibition (from PV to ECS) to promote ECS firing before F. The Hebbian plasticity rule that we use is characterized by a longer time constant of depression than potentiation (and equal maximal amplitudes) and considers the whole history of ECS and F spiking activity (see Materials and Methods and Fig. S2A for more details).

Earlier version: Fig. 2B (bottom) shows an example of the network dynamics with CS present for 2 seconds and US injected after the first second of simulation. CS is described as a Poisson spike train at ∼50Hz, independent of the US input. ECS is active during the whole 2-second interval. CS also affects PV, which is active most of the time; F is active only in the second half of the dynamics when US is also present. All the rhythms generated by the interneurons are apparent in response to CS and simultaneous CS-US inputs, and they are generated by the same mechanisms as in the fear-only network (Fig. 2A). In contrast to the fear-only network, the projection neurons F and ECS are both modulated by the VIP low theta rhythm. This is because the PV increased activity due to CS tends to silence (with the help of SOM) ECS and F during the silent VIP phase at low theta. During the active VIP phase, however, both ECS and F are active, and the simulations show that ECS fires most of the time slightly before F (see Fig. 2B,C and Fig. S1B); this fine timing needed for potentiation is established by the PING rhythm (see Fig. 1C). By contrast, in the first second of the simulation in Fig. 2B (bottom), SOM and PV prevent plasticity by silencing F, which in the absence of US receives only a weak applied current: thus, without concurrent US and CS, STDP cannot occur between ECS and F.

Revised version: Fig. 2B shows an example of the network dynamics with CS present for 2 seconds and US injected after the first second of simulation. ECS is active during the whole 2-second interval. CS also affects PV, which is active most of the time; F is active only in the second half of the dynamics when US is also present. All the rhythms generated by the interneurons are apparent in response to simultaneous effects of the CS-US inputs, and they are generated by the same mechanisms as in the fear-only network (Fig. 2A). In contrast to the fear-only network, the projection neurons F and ECS are both modulated by the VIP low theta rhythm. This is because the PV increased activity due to CS tends to silence (with the help of SOM) ECS and F during the silent VIP phase at low theta. During the active VIP phase, however, both ECS and F are active, and the simulations show that ECS fires most of the time slightly before F (see Fig. 2 B,C and Fig. S2B); this fine timing needed for potentiation is established by the PING rhythm (see Fig. 1C). By contrast, in the first second of the simulation in Fig. 2B (during CS-only), SOM and PV prevent plasticity by silencing F, which in the absence of US receives only a weak applied current.

In the next sections, we will explore the role of each interneuron and its associated rhythm in shaping the network dynamics and allowing the association between CS and fear to be instantiated.

With the depression-dominated plasticity rule, all interneuron types are needed to provide potentiation during fear learning

Earlier version: We now show that, in the example used above, only the network endowed with all the interneurons and their associated rhythms leads to overall potentiation of the conductance from ECS to F in the timeframe used to induce the fear learning in experimental work. (see Discussion for other plasticity rules.) In general, experimental work finds successful learning after one or very few presentations of CS and US (lasting 1.5 or 2 seconds) interspersed with CS-only intervals lasting 30-40 seconds (e.g., see (Davis et al., 2017; Krabbe et al., 2019)). The 40-second interval we consider consists of US and CS active during the entire period: since an initial bout of US is known to produce a long-lasting fear response beyond the offset of the US (Hole and Lorens, 1975), we treat the US as present (e.g., for 1.5 or 2 seconds) followed by US activity induced by memory for the remainder of the 40 seconds.

Revised version: We now show that, in the example used above, only the network endowed with all the interneurons and their associated rhythms leads to overall potentiation of the conductance from ECS to F in the timeframe used to induce the fear learning in experimental work. (See Discussion for other plasticity rules.) In general, experimental work finds successful learning after one or very few presentations of CS and US (lasting 1.5 or 2 seconds) interspersed with CS-only intervals lasting 30-40 seconds (e.g., see (Davis et al., 2017; Krabbe et al., 2019)). The 40-second interval we consider has both ECS and F, as well as VIP and PV interneurons, active during the entire period: an initial bout of US is known to produce a long-lasting fear response beyond the offset of the US (Hole and Lorens, 1975) and to induce the release of neuromodulators. The latter, in particular acetylcholine and dopamine that are known to be released upon US presentation (Harmer and Phillips, 1999; Suzuki et al., 2002; Rajebhosale et al., 2024a), may induce more sustained activity in the ECS, F, VIP, and PV neurons during and after the presentation of US, thus ensuring a concomitant activation of those neurons necessary for STDP to take place (see “Assumptions and predictions of the model” in the Discussion).

Earlier version: Fig. 3A shows the evolution of the average ECS to F AMPA maximal conductance during fear conditioning across 40 network realizations of the full network, as well as from networks lacking VIP, PV, SOM. The average ECS to F AMPA conductance robustly potentiates only in the network containing all three interneurons (Fig. 3A; see Fig. S2 for exceptions outside the normal physiological range). Furthermore, as shown in Fig. 3B, all the full network realizations are “learners”. We define learners as those realizations whose AMPA conductance from ECS to F is higher than 0.12 mS/cm3 at the end of the 40-second interval, which results in the systematic activation of the fear neuron F following most of the ECS spike when only CS is presented. This is consistent with the high rate of successful learning in rodent experiments after one pairing of CS and US, despite inter-individual differences among animals (Schafe et al., 2000). Network realizations differ from one another in the initial state of each neuron involved, and all receive independent Gaussian noise. We note that once the association between CS and fear is acquired, subsequent presentations of CS and US do not weaken or erase it: the interneurons ensure the correct timing and pauses in ECS and F activity, which are conducive for potentiation. Finally, since (Krabbe et al., 2019) reported that a fraction of PV interneurons are affected by US, we have also run the simulations for single neuron network with the PV interneuron affected by US instead of CS. In this case as well, all the network realizations are learners (see Fig. S3).

Revised version: Fig. 3A shows the evolution of the average ECS to F AMPA maximal conductance during fear conditioning across 40 network realizations of the full network, as well as from networks lacking VIP, PV, SOM. The average ECS to F AMPA conductance robustly potentiates only in the network containing all three interneurons (Fig. 3A). Furthermore, as shown in Fig. 3B, all the full network realizations are “learners”. We define learners as those realizations whose AMPA conductance from ECS to F is higher than 0.12 mS/cm3 at the end of the 40-second interval, which results in the systematic activation of the fear neuron F following most of the ECS spikes when only CS is presented. This is consistent with the high rate of successful learning in rodent experiments after one pairing of CS and US, despite inter-individual differences among animals (Schafe et al., 2000). Network realizations differ from one another in the initial state of each neuron involved, and all receive independent Gaussian noise. We note that once the association between CS and fear is acquired, subsequent presentations of CS and US do not weaken or erase it: the interneurons ensure the correct timing and pauses in ECS and F activity, which are conducive for potentiation. Furthermore, since (Krabbe et al., 2019) reported that a fraction of PV interneurons are affected by US, we have also run the simulations for single neuron network with the PV interneuron affected by US instead of CS. In this case as well, all the network realizations are learners (see Fig. S3).

ECS to F conductance across network realizations.

Earlier version: A: Mean (color-coded curves) and standard deviation (color-coded shaded areas) of the AMPA conductance (gAMPA) from ECS to F across 40 network realizations over 40 seconds. Red curve and shaded area represent the mean and standard deviation, respectively, across network realizations endowed with all the interneurons. B: Evolution in time of the AMPA conductance for the 40 full network realizations in A. C: AMPA conductance of 20 network realizations over 45 seconds with F strongly activated by US (2 seconds) and its memory (13 seconds), and ECS active the whole interval because of CS; 19 out of 20 network realizations show potentiation after one trial.

Revised version: A: Mean (color-coded curves) and standard deviation (color-coded shaded areas) of the AMPA conductance (gAMPA) from ECS to F across 40 network realizations over 40 seconds. Red curve and shaded area represent the mean and standard deviation, respectively, across network realizations endowed with all the interneurons. B: Evolution in time of the AMPA conductance for the 40 full network realizations in A. C: AMPA conductance of 20 network realizations over 45 seconds with F strongly activated by US (2 seconds) and its effects (13 seconds), and ECS active the whole interval because of CS; 19 out of 20 network realizations show potentiation after one trial.

Earlier version: We show in Fig. 3C results in which we relax the assumption that US is active over the entire forty seconds by allowing US to be active only during the initial 15 seconds of the fear conditioning paradigm. We find that learning may still occur under this condition. The reason is that, after 15 seconds of CS and US pairing, ECS to F may have potentiated enough that ECS can drive F some of the time (although not all the time). This allows further potentiation to occur in the presence of CS alone. We find that, if the conductance from ECS to F is higher than a threshold value of 0.037 mS/cm3 after the 15 seconds of CS and US presentation, then the network will become a learner after a further thirty seconds of CS alone; those network realizations that did not reach the threshold value are defined here as non-learners. Even with this more realistic assumption, the large majority of network realizations were learners (19 out of 20 networks). This is in agreement with the experimental fear conditioning literature showing that most of the subjects learn the association between CS and fear after only one trial (Schafe et al., 2000).

Revised version: We show in Fig. 3C results in which F and VIP show an increased activity due to the US for only 15 seconds of the fear conditioning paradigm, thus relaxing the assumption of these neurons being active for the entire forty seconds. We find that learning may still occur under this condition. The reason is that, after 15 seconds of enhanced network activity due to CS and US, ECS to F may have potentiated enough that ECS can drive F some of the time (although not all the time). This allows further potentiation to occur in the presence of CS alone. We find that, if the conductance from ECS to F is higher than a threshold value of 0.037 mS/cm3 after 15 seconds, then the network will become a learner after a further thirty seconds of CS alone; those network realizations that did not reach the threshold value are defined here as non-learners. Even with this more restrictive assumption, the large majority of network realizations were learners (19 out of 20 networks). This is in agreement with the experimental fear conditioning literature showing that most of the subjects learn the association between CS and fear after only one trial (Schafe et al., 2000).

Added in the revised version: Mechanisms by which interneurons contribute to potentiation during depression-dominated plasticity

Removed in the revised version: We note that the CS and US inputs modeled as independent Poisson spike trains represent stimuli with no structure. Although we have not explicitly modeled pulsating pips, as common in auditory fear conditioning (e.g., (Stujenske et al., 2014; Krabbe et al., 2019)), we show in Fig. S4 that potentiation can be achieved over a relatively wide range of gamma frequencies. This indicates that overall potentiation is ensured if the gamma frequency transiently increases after the pip.

Removed in the revised version: Physiology of interneuron types is critical to their role in depression-dominated plasticity

Earlier version: The PV cell is necessary to induce the correct pre-post timing between ECS and F needed for long-term potentiation of the ECS to F conductance. In our model, PV has reciprocal connections with F and provides lateral inhibition to ECS. Since the lateral inhibition is weaker than the feedback inhibition, PV tends to bias ECS to fire before F. This creates the fine timing needed for the depression-dominated rule to instantiate plasticity. If we used the classical Hebbian plasticity rule (Bi and Poo, 2001) with gamma frequency inputs, this fine timing would not be needed and ECS to F would potentiate over most of the gamma cycle, and thus we would expect random timing to lead to potentiation (Fig. S5). In this case, no interneurons are needed (See Discussion “Synaptic plasticity in our model” for the potential necessity of the depression-dominated rule).

Revised version: The PV cell is necessary to induce the correct pre-post timing between ECS and F needed for long-term potentiation of the ECS to F conductance. In our model, PV has reciprocal connections with F and provides lateral inhibition to ECS. Since the lateral inhibition is weaker than the feedback inhibition, PV tends to bias ECS to fire before F. This creates the fine timing needed for the depression-dominated rule to instantiate plasticity. If we used the classical Hebbian plasticity rule (Bi and Poo, 2001) with gamma frequency inputs, this fine timing would not be needed and ECS to F would potentiate over most of the gamma cycle, and thus we would expect random timing between ECS and F to lead to potentiation (Fig. S4). In this case, no interneurons are needed (See Discussion “Synaptic plasticity in our model” for the potential necessity of the depression-dominated rule).

Earlier version: In this network configuration, the pre-post timing for ECS and F is repeated robustly over time due to coordinated gamma oscillations (PING, as shown in Fig. 4A, Fig. 1C, and Fig. 2B,C) arising through the reciprocal interactions between F and PV (Feng et al., 2019). PING can arise only when PV is in a sufficiently low excitation regime such that F can control PV activity (Börgers et al., 2005), as in Fig. 4A. However, despite PING establishing the correct spike-timing for potentiation, setting PV in a low excitation regime in which F-PV PING occurs continuously does not lead to potentiation (Fig. 4A, Fig. S1C): the depression-dominated rule leads to depression rather than potentiation unless the PING is periodically interrupted. During the pauses, made possible only in the full network by the presence of VIP and SOM, the history-dependent build-up of depression decays back to baseline, allowing potentiation to occur on the next ECS/F active phase. (The detailed mechanism of how this happens is in the Supplementary Information, including Fig. S1). Thus, a network without the other interneuron types cannot lead to potentiation.

Revised version: In the PV and projection cell network configuration, the pre-post timing for ECS and F is repeated robustly over time due to coordinated gamma oscillations (PING, as shown in Fig. 4A, Fig. 1C) arising through the reciprocal interactions between F and PV (Feng et al., 2019). PING can arise only when PV is in a sufficiently low excitation regime such that F can control PV activity (Börgers et al., 2005), as in Fig. 4A. However, although such a low excitation regime establishes the correct fine timing for potentiation, it is not sufficient to lead to potentiation (Fig. 4A, Fig. S2C): the depression-dominated rule leads to depression rather than potentiation unless the PING is periodically interrupted. During the pauses, made possible only in the full network by the presence of VIP and SOM, the history-dependent build-up of depression decays back to baseline, allowing potentiation to occur on the next ECS/F active phase. (The detailed mechanism of how this happens is in the Supplementary Information, including Fig. S2). Thus, a network without the other interneuron types cannot lead to potentiation. Though a low excitation level for a PV cell is necessary to produce a PING, a higher excitation level is necessary to produce a pause in the ECS and F. This higher excitation level is consistent with the experimental literature showing a strong activation of PV after the onset of CS (Wolff et al., 2014). The higher excitation happens when the VIP cell is silent, whereas a low excitation level is achieved when the VIP cell fires and partially inhibits the PV cell (Fig. 4B, Fig. S2D). The interruption in the ECS and F activity requires the participation of another interneuron, the SOM cell (Figs. 2B, S2): the pauses in inhibition from the VIP periodically interrupt ECS and F firing by releasing PV and SOM from inhibition and thus indirectly silencing ECS and F. Without these pauses, depression dominates (see SI section “ECS and F activity patterns determine overall potentiation or depression”).

PV-only network and VIP-PV network lead to depression.

Earlier version: A: Left, network with PV as the only interneuron. PV cell is at an excitation level that supports PING. Middle, dynamics of PV and F that reciprocally interact and generate PING. Right, PV and F entrain in PING (top), ECS and F activity (middle), and ECS to F conductance (bottom). B: Left, network with both PV and VIP. Right, network dynamics (top, middle) followed by the evolution in time of the ECS to F AMPA conductance (bottom). The detailed mechanism behind the evolution of AMPA conductance in panels A and B is in the Supplementary Information, including Fig. S1.

Revised version: A: Left, network with PV as the only interneuron. PV cell is at an excitation level that supports PING. Middle, dynamics of PV and F that reciprocally interact and generate PING. Right, PV and F entrain in PING (top), ECS and F activity (middle), and ECS to F conductance (bottom). B: Left, network with both PV and VIP. Right, network dynamics (top, middle) followed by the evolution in time of the ECS to F AMPA conductance (bottom). The detailed mechanism behind the evolution of AMPA conductance in panels A and B is in the Supplementary Information, including Fig. S2.

Removed in the revised version: In our simulations (except for the one in Fig. 4A), we choose to set PV cells at a high excitation level due to a strong input from CS in order to account for: (i) the experimental literature showing a strong activation of PV after the onset of CS (Wolff et al., 2014); (ii) many other inputs that the PV cell receive besides CS. We show in Fig. S2A that increasing the PV excitation by adding other inputs to PV cells leads to similar results; all interneurons are necessary for ECS to F potentiation under these circumstances. If CS activation of PV is weakened, we find potentiation can occur in the absence of VIP and SOM (Fig. S2B). However, we consider this outside the physiological range.

Removed in the revised version: In the context of high PV excitation, we find that the VIP cell plays two critical roles in enabling potentiation to occur via PING. First, it reduces the excitation of the PV cell during its active low-theta phase, enabling the PV cell to participate in PING (Fig. 4B, Fig. S1D). Second, it provides the periodic interruptions in ECS and F firing necessary for potentiation during its silent low-theta phase. Without these pauses, depression dominates (see SI section “ECS and F activity patterns determine overall potentiation or depression”). This interruption requires the participation of another interneuron, the SOM cell (Figs. 2B, 4B, S1): the pauses in inhibition from the VIP periodically interrupt ECS and F firing by releasing PV and SOM from inhibition and thus indirectly silencing ECS and F.

Network with multiple heterogeneous neurons can establish the association between CS and fear

Earlier version: To test the robustness of our single-cell results to heterogeneity, we expand our BLA network to include three cells of each interneuron subtype and ten of each excitatory projection neuron (Fig. 5A). Each neuron has independent noise and cellular parameters (see Materials and Methods for details). We find that the network very robustly produces potentiation between the ECS and F receiving CS and US, respectively, during fear training.

Revised version: To test the robustness of our single cell results to heterogeneity, we expand our BLA network to include three cells of each interneuron subtype and ten of each excitatory projection neuron (Fig. 5A). Each neuron has independent noise and cellular parameters (see Materials and Methods for details). We find that the network very robustly produces potentiation between the ECS and F affected by CS and US, respectively, during fear training.

Heterogeneous BLA fear network is capable of establishing the association between CS and fear.

Earlier version: A: whole BLA network with multiple and heterogeneous neurons. B: Dynamics in the first 5000 ms of each of the neurons in the BLA network. C: Dynamics of ECS to F conductance over 5000 ms shaped by the activity in B. D: Left, mean and standard deviation across 40 network realizations of the ECS to F conductance for the full (red), no VIP (green), no SOM (purple), no PV (black), no SOM and PV (magenta) networks. The green, purple, back, and magenta curves are superimposed on each other. Right, dynamics of all the 40 full network realizations over 40 seconds.

Revised version: A: whole BLA network with multiple and heterogeneous neurons. B: Dynamics in the first 5000 ms after the onset of US and CS of each of the neurons in the BLA network. C: Dynamics of ECS to F conductance over 5000 ms shaped by the activity in B. D: Left, mean and standard deviation across 40 network realizations of the ECS to F conductance for the full (red), no VIP (green), no SOM (purple), no PV (black), no SOM and PV (magenta) networks. The green, purple, back, and magenta curves are superimposed on each other. Right, dynamics of all the 40 full network realizations over 40 seconds.

Earlier version: Fig. 5B shows an example of the network dynamics during fear conditioning with simultaneous CS and US inputs. As previously presented for the single neuron network (Figs. 2-3), interneurons are crucial in conferring the correct pre-post spike timing to ECS and F. We assume all the VIP interneurons receive the same US; hence, the VIP neurons tend to approximately synchronize at low theta in response to US input, allowing a window for potentiation of ECS to F conductance (Fig. 5B), as we have seen in the single cell model. The potentiation from ECS to F is specifically for the excitatory projection cells receiving CS and US, respectively (Fig. 5C). The ECS cells not receiving CS are inhibited by ongoing PV activity during the disinhibition window (Fig. 5B); they are constructed to be firing at 11 Hz in the absence of any connections from other cells. The lack of activity in those cells during fear conditioning implies that there is no plasticity from those ECS cells to the active F. Those cells are included for the calculation of the LFP (see below in “Increased low-theta frequency is a biomarker of fear learning”). This larger network corroborates the results obtained for the single neuron network: only the realizations of the full network learn the association between CS and fear (Fig. 5D, left) and all those network realizations become learners in less than forty seconds (Fig. 5D, right). Similarly, there is a striking failure of plasticity if any interneuron is removed from the network; even partial plasticity does not arise.

Revised version: Fig. 5B shows an example of the network dynamics during fear conditioning. As previously presented for the single neuron network (Figs. 2-3), interneurons are crucial in conferring the correct pre-post spike timing to ECS and F. We assume all the VIP interneurons receive the same US; hence, the VIP neurons tend to approximately synchronize at low theta, allowing a coordinated disinhibition window for potentiation of ECS to F conductance (Fig. 5B), as we have seen in the single cell model. The potentiation from ECS to F is specifically for the excitatory projection cells affected by CS and US, respectively (Fig. 5C). The ECS cells not receiving CS are inhibited by ongoing PV activity during the disinhibition window (Fig. 5B); they are constructed to be firing at 11 Hz in the absence of any connections from other cells. The lack of activity in those cells during fear conditioning implies that there is no plasticity from those ECS cells to the active F. Those cells are included for the calculation of the LFP (see below in “Increased low-theta frequency is a biomarker of fear learning”). This larger network corroborates the results obtained for the single neuron network: only the realizations of the full network learn the association between CS and fear (Fig. 5D, left) and all those network realizations become learners in less than forty seconds (Fig. 5D, right). Similarly, there is a striking failure of plasticity if any interneuron type is removed from the network; even partial plasticity does not arise.

Increased low-theta frequency is a biomarker of fear learning

Earlier version: In this section, we explore how plasticity in the fear circuit affects the network dynamics, comparing after fear conditioning to before. We first show that fear conditioning leads to an increase in low theta frequency power of the network spiking activity compared to the pre-conditioned level (Fig. 6 A,B); there is no change in the high theta power. We also show that the LFP, modeled as the linear sum of all the AMPA, GABA, NaP-, D-, and H-currents in the network, similarly reveals a low theta power increase and no significant variation in the high theta power (Fig. 6 C,D,E). These results reproduce the experimental findings in (Davis et al., 2017), and Fig. 6 F,G show that the low theta increase is due to added excitation provided by the new learned pathway. The additional unresponsive ECS and F cells in the network were included to ensure we had not biased the LFP towards excitation. Nevertheless, although both the AMPA and GABA currents contribute to the power increase in the low theta frequency range (Fig. 6F), the AMPA currents show a dramatic power increase relative to the baseline (the average power ratio of AMPA and GABA post-vs pre-conditioning across 20 network realizations is 3*103 and 4.6, respectively). This points to the AMPA currents as the major contributor to the low theta power increase. Specifically, the newly potentiated AMPA synapse from ECS to F ensures F is active after fear conditioning, thus generating strong currents in the PV cells to which it has strong connections (Fig. 6G). Finally, the increase in power is in the low theta range because ECS and F are allowed to spike only during the active phase of the low theta spiking VIP neurons. We have also explored another proxy for the LFP (see Supplementary Information and Fig. S6).

Revised version: We find that fear conditioning leads to an increase in low theta frequency power of the network spiking activity compared to the pre-conditioned level (Fig. 6 A,B); there is no change in the high theta power. We also find that the LFP, modeled as the linear sum of all the AMPA, GABA, NaP-, D-, and H-currents in the network, similarly reveals a low theta power increase when considering the peak of the low theta power, and no significant variation in the high theta power again when considering the peak of the high theta power (Fig. 6 C,D,E). These results are consistent with the experimental findings in (Davis et al., 2017). Specifically, the newly potentiated AMPA synapse from ECS to F ensures F is active after fear conditioning, thus generating strong currents in the PV cells to which it has strong connections (Fig. 6F). It is the AMPA currents to the PV interneurons that are directly responsible for the low theta increase; it is the newly potentiated ECS to F synapse that paces the AMPA currents in the PV interneurons to go at low theta. Thus, the low theta increase is due to added excitation provided by the new learned pathway.

Heterogeneous network dynamics and spectral properties pre versus post fear conditioning for network realizations in learners and non-learners.

Earlier version: All power spectra are represented as mean and standard deviation across 20 network realizations. A: Dynamic of BLA heterogeneous networks pre (left) and post (learner, middle; non-learner, right) fear conditioning. B: Power spectra of network spiking activity before fear conditioning (blue) and after successful (purple) and non-successful fear conditioning (orange); top, right: inset between 2 and 6 Hz. Blue and orange curves closely overlap. C: power spectra of the LFP proxy (linear sum of AMPA, GABA, D-, NaP-, and H-currents); all the details as in B. D, E: 25th, 50th and 75th percentiles of low theta power in the 2.5-4 Hz range (D) and high theta power in the 12-14 Hz range (E) in 20 network realizations before and after (in both learners and non-learners) fear conditioning. ***: p-value < 10−5; n.s.: non-significant difference. F: Power spectra mean and standard deviation of the LFP signals derived from only AMPA currents (red curve), GABA currents (green curve), D-current, NaP-current and H-current (light blue curve). AMPA currents are generated by the interactions from ECS to F, F to VIPs, and F to PVs. VIP cells contribute to the D-current and SOM cells to H-current and NaP-current (see the result’s section “Rhythms in the BLA can be produced by interneurons” for a description of these currents). G: Power spectra of AMPA currents from ECS to F (black curve) and from F to PV interneurons (red curve).

Revised version: All power spectra are represented as mean and standard deviation across 20 network realizations. A: Dynamic of BLA heterogeneous networks pre (left) and post (learner, middle; non-learner, right) fear conditioning. B: Power spectra of network spiking activity before fear conditioning (blue) and after successful (purple) and non-successful fear conditioning (orange); top, right: inset between 2 and 6 Hz. Blue and orange curves closely overlap. C: Power spectra of the LFP proxy (linear sum of AMPA, GABA, D-, NaP-, and H-currents); all the details as in B. D, E: 25th, 50th and 75th percentiles of LFP low theta power in the 2.5-4 Hz range where the peaks of power exist (D) and high theta power in the 12-14 Hz range again where the peaks of power exist (E) in 20 network realizations before and after (in both learners and non-learners) fear conditioning. ***: p-value < 0.001; n.s.: non-significant difference. F: Power spectra of AMPA currents from ECS to F (black curve) and from F to PV interneurons (red curve). G: Power spectra mean and standard deviation of the LFP signals derived from only AMPA currents (red curve), GABA currents (green curve), D-current, NaP-current and H-current (light blue curve). AMPA currents are generated by the interactions from ECS to F, F to VIPs, and F to PVs. VIP cells contribute to the D-current and SOM cells to H-current and NaP-current (see the result’s section “Rhythms in the BLA can be produced by interneurons” for a description of these currents).

Added in the revised version: We find that the AMPA currents are the major contributor to the low theta increase. Although both the AMPA and GABA currents contribute to the power increase in the low theta frequency range (Fig. 6G), the AMPA currents show a dramatic power increase relative to the baseline (the average power ratio of AMPA and GABA post-vs pre-conditioning across 20 network realizations is 3*103 and 4.6, respectively). Thus, the AMPA currents are the major contributors to the low theta power increase. As a further constrain, the additional unresponsive ECS and F cells in the network were included to ensure we had not biased the LFP towards excitation. Finally, the increase in power is in the low theta range because ECS and F are allowed to spike only during the active phase of the low theta spiking VIP neurons. We have also explored another proxy for the LFP (see Supplementary Information and Fig. S5).

Earlier version: Although the experimental results in (Davis et al., 2017) show an increase in low theta after fear learning, they excluded non-learners from their analysis, and thus, it is unclear from their results if low theta can be used as a biomarker of fear learning. To address this question, we looked at the power spectra after fear conditioning in learning versus non-learning networks. To have an adequate number of non-learners, we ran 60 network simulations for 10 seconds and chose those whose conductance from ECS to F remained lower than 0.037 mS/cm3 (i.e., non-learners). Notably, the low theta power increase is completely absent after fear conditioning in those network realizations that display no signs of learning. This suggests that the low theta power change is a biomarker of successful fear conditioning: it occurs when there is learning and does not occur when there is no learning.

Revised version: Although the experimental results in (Davis et al., 2017) show an increase in low theta after fear learning, they excluded non-learners from their analysis, and thus, it is unclear from their results if low theta can be used as a biomarker of fear learning. To address this question, we looked at the power spectra after fear conditioning in learning versus non-learning networks. To have an adequate number of non-learners, we ran 60 network simulations for 10 seconds and chose 20 of them whose conductance from ECS to F remained lower than 0.037 mS/cm3 (i.e., non-learners). Notably, the low theta power increase is completely absent after fear conditioning in those network realizations that display no signs of learning. This suggests that the low theta power change is a biomarker of successful fear conditioning: it occurs when there is learning and does not occur when there is no learning.

Discussion

Overview

Our study suggests that amygdalar rhythms play a crucial role in plasticity during fear conditioning. Prominent rhythms found in the BLA during fear conditioning include low theta (3-6 Hz), high theta (6-12 Hz), and gamma (>30 Hz) (Seidenbecher et al., 2003; Courtin et al., 2014b; Stujenske et al., 2014; Davis et al., 2017). Experimental work in rodents shows that the BLA undergoes more oscillatory firing at low theta frequency and gamma following fear conditioning, while high theta frequency remains unchanged compared to before fear conditioning (Courtin et al., 2014b; Davis et al., 2017).

To examine the origin of these rhythmic changes and their functional role in fear conditioning, we implement a biophysically detailed model of the BLA. We show that vasoactive intestinal peptide-expressing (VIP) interneurons, parvalbumin-expressing (PV) interneurons, and somatostatin-expressing (SOM) interneurons in the BLA may be centrally involved in producing the experimentally measured rhythms based on the biophysical properties of these interneurons. More specifically: the gamma oscillation is associated with PV cell interaction with excitatory cells to produce PING; VIP cells produce low theta due to their intrinsic D-current; SOM cells help produce the high theta rhythm due to persistent sodium current (NaP-current) and H-currents. Moreover, we show that the rhythmic dynamics produced by VIP and PV cells both play a crucial role in the instantiation of plasticity during associative learning by promoting the formation of a dedicated fear circuit shaped through spike-timing-dependence, and the removal of any of the interneuron types from the circuit leads to failure of the plasticity needed for associative fear learning. We note that the presence of SOM cells is crucial for plasticity in our model since they help to produce the necessary pauses in the excitatory projection cell activity. The BLA SOM cells do not necessarily have to be the only source of the high theta observed in the BLA during fear learning; the high theta detected in the LFP of the BLA also originates from the prefrontal cortex and/or the hippocampus (Stujenske et al., 2014, 2022). Finally, we replicate the experimental increase in the low theta after successful training (Davis et al., 2017) and determine that this is a biomarker of fear learning, i.e., the increase in low-theta power does not appear in the absence of learning.

Synaptic plasticity in our model

Earlier version: Synaptic plasticity is the mechanism underlying the association between neurons that respond to the neutral stimulus CS (ECS) and those that respond to fear (F), which instantiates the acquisition and expression of fear behavior. One form of experimentally observed long-term synaptic plasticity is spike-timing-dependent plasticity (STDP), which defines the amount of potentiation and depression for each pair of pre- and postsynaptic neuron spikes as a function of their relative timing (Bi and Poo, 2001; Caporale and Dan, 2008). There are many STDP rules in the literature (Abbott and Nelson, 2000; Feldman, 2012) but none we are aware of specifically for the amygdala. The depression-dominated Hebbian rule, which we use in this study, provides each interneuron subtype with a specific role in allowing the instantiation of fear learning. With this rule, depression is overall stronger than potentiation, unless there is fine timing among pyramidal cells in the gamma and low theta cycles.

Revised version: Synaptic plasticity is the mechanism underlying the association between neurons that respond to the neutral stimulus CS (ECS) and those that respond to fear (F), which instantiates the acquisition and expression of fear behavior. One form of experimentally observed long-term synaptic plasticity is spike-timing-dependent plasticity (STDP), which defines the amount of potentiation and depression for each pair of pre- and postsynaptic neuron spikes as a function of their relative timing (Bi and Poo, 2001; Caporale and Dan, 2008). All forms of STDP require that there be an overlap in the firing of the pre- and post-synaptic cells. In some fear learning paradigms, the US and the CS do not overlap. We address this below under “Assumptions and predictions of the model”, showing how the effects of US and CS on the spiking of the relevant neurons can overlap even in the absence of overlap of US and CS.

Earlier version: With the depression-dominated rule, each interneuron plays a role in setting up this fine timing in the gamma and low theta cycles needed for appropriate plasticity. PV interneurons are important in creating a gamma rhythm with F when the latter is activated by the unconditioned stimulus (US). This gamma rhythm plays a central role in the fine timing between ECS and F. As mentioned earlier, the critical requirements for plasticity are a pause of VIP firing within each low-theta cycle and fine timing of the ECS (pre) and F (post) neurons in the active phase of the cycle. The pause in VIP cell firing allows for activation of SOM cells, which helps to inhibit the ECS and F cells. Without the pause, the ECS and F cells continue to fire. It is a consequence of the depression-dominated rule that this would lead to depression; this happens because the depression does not relax to zero between gamma-frequency spikes, and hence continues to build up. With a Hebbian plasticity rule characterized by a lower amplitude for depression than for potentiation, potentiation would occur at most of the phases of gamma and thus fine timing would not be needed (Fig. S5). The depression-dominated rule allows for regulation of the plasticity by modulation of all the kinds of cells that are known to be involved in fear learning as well as providing reasons for the known involvement of rhythms.

Revised version: There are many STDP rules in the literature (Abbott and Nelson, 2000; Feldman, 2012) but none we are aware of specifically for the amygdala. With the depression-dominated rule used in this study, depression is overall stronger than potentiation, unless there is fine timing among pyramidal cells in the gamma and low theta cycles. With this rule, each interneuron plays a role in setting up this fine timing in the gamma and low theta cycles needed for appropriate plasticity. PV interneurons are important in creating a gamma rhythm with F when the latter is activated by the unconditioned stimulus (US). This gamma rhythm plays a central role in the fine timing between ECS and F. As mentioned in the Result section “Mechanisms by which interneurons contribute to potentiation during depression-dominated plasticity”, the critical requirements for plasticity are a pause of VIP firing within each low-theta cycle and fine timing of the ECS (pre) and F (post) neurons in the active phase of the cycle. The pause in VIP cell firing allows for activation of SOM cells, which helps to inhibit the ECS and F cells. Without the pause, the ECS and F cells continue to fire. It is a consequence of the depression-dominated rule that this would lead to depression; this happens because the depression does not relax to zero between gamma-frequency spikes, and hence continues to build up. With a Hebbian plasticity rule characterized by a lower amplitude for depression than for potentiation, potentiation would occur at most of the phases of gamma and thus fine timing would not be needed (Fig. S4). The depression-dominated rule allows for regulation of the plasticity by modulation of all the kinds of cells that are known to be involved in fear learning as well as providing reasons for the known involvement of rhythms.

Earlier version: With the classical Hebbian plasticity rule, we show that learning can occur without the involvement of the VIP and SOM cells. Although fear learning can occur without the depression-dominated rule, we hypothesize that the latter is necessary for other aspects of fear learning and regulation. Generalization of learning can be pathological, and we hypothesize that the modulation created by the involvement of VIP and SOM interneurons is normally used to prevent such overgeneralization. However, in some circumstances, it may be desirable to account for many possible threats, and then a classical Hebbian plasticity rule could be useful. We note that the involvement or not of the VIP-SOM circuit has been implicated when there are multiple strategies for solving a task (Piet et al., 2024). In our situation, the nature of the task (including reward structure) may determine whether the learning rule is depression-dominated and therefore whether the VIP-SOM circuit plays an important role.

Revised version: With the classical Hebbian plasticity rule, we suggest that learning can occur without the involvement of the VIP and SOM cells. Although fear learning can occur without the depression-dominated rule, we hypothesize that the latter is necessary for other aspects of fear learning and regulation. Generalization of learning can be pathological, and we hypothesize that the modulation created by the involvement of VIP and SOM interneurons is normally used to prevent such overgeneralization. However, in some circumstances, it may be desirable to account for many possible threats, and then a classical Hebbian plasticity rule could be useful. We note that the involvement or not of the VIP-SOM circuit has been implicated when there are multiple strategies for solving a task (Piet et al., 2024). In our situation, the nature of the task (including reward structure) may determine whether the learning rule is depression-dominated and therefore whether the VIP-SOM circuit plays an important role.

Earlier version: VIP cells may be important in establishing the kind of depression-dominated rule we are using in this model for the synapse between ECS and F. VIP cells are known to corelease the peptide VIP along with GABA (Bayraktar et al., 1997). The amount of the release is related to the amount of high frequency firing, i.e., the duty cycle of the gamma burst in each low theta cycle (Agoston et al., 1988; Agoston and Lisziewicz, 1989). The peptide VIP can act on second messenger pathways to inhibit potentiation; the pathways are complex and not fully understood, but in hippocampus they involve GABA transmission, NMDA activation, and CaMKII (Caulino-Rocha et al., 2022). The relevant VIP receptor VPAC1 is known to exist in the amygdala (Joo et al., 2005; Boucher et al., 2021). Thus, by inhibiting potentiation, VIP may change a Hebbian plasticity rule to a depression-dominated rule. (See (Cuhna-Reis et al., 2020) for details about VIP and plasticity). We hypothesize that SST peptide is critical in fear extinction and preventing fear generalization, but not to initial fear learning. Also, the CCK peptide has been proposed to promote the switch between fear and safety states after fear extinction (Krabbe et al., 2018).

Revised version: VIP cells may be important in establishing the kind of depression-dominated rule we are using in this model for the synapse between ECS and F. VIP cells are known to corelease the peptide VIP along with GABA (Bayraktar et al., 1997). The amount of the release is related to the amount of high frequency firing, i.e., the duty cycle of the gamma burst in each low theta cycle (Agoston et al., 1988; Agoston and Lisziewicz, 1989). The peptide VIP can act on second messenger pathways to inhibit potentiation; the pathways are complex and not fully understood, but in hippocampus they involve GABA transmission, NMDA activation, and CaMKII (Caulino-Rocha et al., 2022). The relevant VIP receptor VPAC1 is known to exist in the amygdala (Joo et al., 2005; Boucher et al., 2021). Thus, by inhibiting potentiation, VIP may change a Hebbian plasticity rule to a depression-dominated rule. (See (Cunha-Reis and Caulino-Rocha, 2020) for details about VIP and plasticity). Since high-theta is increased during fear extinction, but not learning, and our model suggests an amygdalar source of high-theta is the SOM cells, we hypothesize that SST peptide may play some role in aiding the formation of networks involved in fear extinction. Also, the CCK peptide has been proposed to promote the switch between fear and safety states after fear extinction (Krabbe et al., 2018).

Involvement of other brain structures

Earlier version: Studies using fear conditioning as a model of associative learning reveal that learning and expression of fear are not limited to the amygdala but involve a distributed network including the amygdala, the medial prefrontal cortex, and the hippocampus (Seidenbecher et al., 2003; Bocchio and Capogna, 2014; Courtin et al., 2014a; Stujenske et al., 2014; Tovote et al., 2015; Karalis et al., 2016; Chen et al., 2021). In our model, the fear-encoding neuron F fires for at least 13 seconds due to the US (first ∼2 seconds), followed by the fear memory that could excite F in a similar way to US. This is suggested by the observation that during fear conditioning, animals are hyperactive for about 10 seconds after US presentation and then freeze while they are maintained in the box in which they received the shock (Hole and Lorens, 1975). The medial prefrontal cortex and the hippocampus may provide the substrates for the continued firing of F needed for ECS to F potentiation after the 2-second US stimulation.

Revised version: Studies using fear conditioning as a model of associative learning reveal that learning and expression of fear are not limited to the amygdala but involve a distributed network including the amygdala, the medial prefrontal cortex, and the hippocampus (Seidenbecher et al., 2003; Bocchio and Capogna, 2014; Courtin et al., 2014a; Stujenske et al., 2014; Tovote et al., 2015; Karalis et al., 2016; Chen et al., 2021). In our model, the excitatory projection neurons and VIP and PV interneurons show sustained activity during and after the US presentation, thus allowing potentiation through STDP to take place. The medial prefrontal cortex and/or the hippocampus may provide the substrates for the continued firing of the BLA neurons after the 2-second US stimulation. We also discuss below that this network sustained activity may originate from neuromodulator release induced by US (see section “Assumptions and predictions of the model” in the Discussion).

Earlier version: Other brain structures may be involved in later stages of fear responsiveness, such as fear extinction and prevention of generalization. It has been reported that the prelimbic cortex (PL) modulates the BLA SOM cells during fear retrieval, and the latter cells are crucial to discriminate non-threatening cues when desynchronized by the PL inputs (Stujenske et al., 2022). Brain structures such as the prefrontal cortex and hippocampus have been documented to play a crucial role also in fear extinction, the paradigm following fear conditioning aimed at decrementing the conditioned fearful response through repeated presentations of the CS alone. As reported by several studies, fear extinction suppresses the fear memory through the acquisition of a distinct memory, instead of through the erasure of the fear memory itself (Harris et al., 2000; Bouton, 2002; Trouche et al., 2013; Thompson et al., 2018). Davis et al., 2017 found a high theta rhythm following fear extinction that was associated with the suppression of threat in rodents. Our model can be extended to include structures in the prefrontal cortex and the hippocampus to further investigate the role of rhythms in the context of discrimination of non-threatening cues and extinction. We hypothesize that a different population of PV interneurons plays a crucial role in mediating competition between fearful memories, associated with a low theta rhythm, and safety memories, associated with a high theta rhythm; supporting experimental evidence is in (Lucas et al., 2016; Davis et al., 2017; Chen et al., 2022).

Revised version: Other brain structures may be involved in other aspects of fear responsiveness, such as fear extinction and prevention of generalization. It has been reported that the prelimbic cortex (PL) modulates the BLA SOM cells during fear retrieval, and the latter cells are crucial to discriminate non-threatening cues when desynchronized by the PL inputs (Stujenske et al., 2022). Also, brain structures such as the prefrontal cortex and hippocampus have been documented to play a crucial role in fear extinction, the paradigm following fear conditioning aimed at decrementing the conditioned fearful response through repeated presentations of the CS alone. As reported by several studies, fear extinction suppresses the fear memory through the acquisition of a distinct memory, instead of through the erasure of the fear memory itself (Harris et al., 2000; Bouton, 2002; Trouche et al., 2013; Thompson et al., 2018). Davis et al., 2017 found a high theta rhythm following fear extinction that was associated with the suppression of threat in rodents. Our model can be extended to include structures in the prefrontal cortex and the hippocampus to further investigate the role of rhythms in the context of discrimination of non-threatening cues and extinction. We hypothesize that a different population of PV interneurons plays a crucial role in mediating competition between fearful memories, associated with a low theta rhythm, and safety memories, associated with a high theta rhythm; supporting experimental evidence is in (Lucas et al., 2016; Davis et al., 2017; Chen et al., 2022).

Added in the revised version: Where the rhythms originate, and by what mechanisms

Added in the revised version: A recent experimental paper (Antonoudiou et al., 2022) suggests that the BLA can intrinsically generate theta oscillations (3-12 Hz) detectable by LFP recordings when inhibition is totally removed due to gabazine application. They draw this conclusion in mice by removing the hippocampus, which can volume conduct to BLA, and noticing that other nearby brain structures did not display any oscillatory activity. In our model, we note that when inhibition is removed, both AMPA and intrinsic currents contribute to the network dynamics and the LFP. Thus, interneurons with their specific intrinsic currents (i.e., D-current in the VIP interneurons, and NaP- and H-currents in SOM interneurons) can indeed affect the model LFP and support the generation of theta and gamma rhythms (Fig. 6G).

Added in the revised version: Another slice study, (Bratsch-Prince et al., 2024), shows that BLA is intrinsically capable of producing a low theta rhythm with ACh stimulation and without needing external glutamate input. ACh is produced in vivo by the basal forebrain in response to US (Rajebhosale et al., 2024b). Although we did not explicitly include the BF and ACh modulation of BLA in our model, we implicitly include the effect of ACh in BLA by increasing the activity of the VIP cells, which then produce the low theta rhythm. Indeed, low theta in the BLA is known to depend on the muscarinic activation of CCK interneurons, a group of interneurons that overlaps with the class of VIP neurons in our model (Mascagni and McDonald, 2003; Krabbe et al., 2018).

Added in the revised version: Although the BLA can produce these rhythms, this does not rule out that other brain structures also produce the same rhythms through different mechanisms, and these can be transmitted to the BLA. Specifically, it is known that the olfactory bulb produces and transmits the respiratory-related low theta (4 Hz) oscillations to the dorsomedial prefrontal cortex, where it organizes neural activity (Bagur et al., 2021). Thus, the respiratory-related low theta may be captured by BLA LFP because of volume conduction or through BLA extensive communications with the prefrontal cortex. Furthermore, high theta oscillations are known to be produced by the hippocampus during various brain functions and behavioral states, including during spatial exploration (Vanderwolf, 1969) and memory formation/retrieval (Raghavachari et al., 2001), which are both involved in fear conditioning. Similarly to the low theta rhythm, the hippocampal high theta can manifest in the BLA. It remains to understand how these other rhythms may interact with the ones described in our paper. However, we emphasize that there is also evidence (as discussed above) that these rhythms arise within the BLA.

Added in the revised version: Assumptions and predictions of the model

Added in the revised version: The interneuron descriptions in the model were constrained by the electrophysiological properties reported in response to hyperpolarizing currents (Sosulina et al., 2010). Specifically, we modeled the three subtypes of VIP, SOM, and PV interneurons displaying bursting behavior, regular spiking with early spike-frequency adaptation, and regular spiking without spike-frequency adaptation, respectively. Focusing on VIP interneurons, we were able to model the bursting behavior by including the D-type potassium current. This current is thought to exist in the VIP interneurons in the cortex (Porter et al., 1998), but whether this current is also found in the BLA is still unknown. We modeled the SOM interneurons with NaP- and H-currents, similar to the OLM cells in the hippocampus. Due to these currents, the VIP and SOM cells are able to produce low- and high-theta oscillations, respectively. The presence of these currents, and the neurons’ ability to exhibit oscillations in the theta range during fear conditioning and at baseline in BLA, which are assumptions of our model, should be tested experimentally. Our model predicts that blockade of D-current in VIP interneurons (or silencing VIP interneurons) will both diminish low theta and prevent fear learning. Finally, the model assumes the absence of significantly strong connections from the excitatory projection cells ECS to PV interneurons, unlike the ones from F to PV. Including those synapses would alter the PING rhythm created by the interactions between F and PV, which is crucial for fine timing between ECS and F needed for LTP.

Added in the revised version: Our model, which is a first effort towards a biophysically detailed description of the BLA rhythms and their functions, does not include the neuron morphology, many other cell types, conductances, and connections that are known to exist in the BLA; models such as ours are often called “minimal models” and constitute most biologically detailed models. For example, although there is considerable variability in the activity patterns of both VIP cells and SOM cells (Sosulina et al., 2010; Guthman et al., 2020; Ünal et al., 2020; Vereczki et al., 2021), our focus was specifically on those subtypes that generate critical rhythms within the BLA. Such minimal models are used to maximize the insight that can be gained by omitting details whose influence on the answers to the questions addressed in the model are believed not to be qualitatively important. We note that the absence of these omitted features constitutes hypotheses of the model: we hypothesize that the absence of these features does not materially affect the conclusions of the model about the questions we are investigating. Of course, such hypotheses can be refuted by further work showing the importance of some omitted features for these questions and may be critical for other questions. Our results hold when there is some degree of heterogeneity of cells of the same type, showing that homogeneity is not a necessary condition.

Added in the revised version: Our study suggests that all the interneurons modeled are necessary for associative learning provided that the STDP rule is depression-dominated. This prediction could be tested experimentally by selectively silencing each interneuron subtype in the BLA: if the associative learning is hampered by silencing any of the interneuron subtypes, this supports our study. Finally, the model prediction could be tested indirectly by acquiring more information about the plasticity rule involved in the BLA during associative learning. We found that all the interneurons are necessary to establish fear learning only in the case of a depression-dominated rule. This rule ensures that fine timing and pauses are always required for potentiation: interneurons provide both fine timing and pauses to pyramidal cells, making them crucial components of the fear circuit.

Added in the revised version: Finally, our model requires the effect of the CS and US inputs on the BLA neuron activity to overlap in time in order to instantiate fear learning through STDP. Such a hypothesis, that learning uses spike-timing-dependent plasticity, is common in the modeling literature (Bi and Poo, 2001; Caporale and Dan, 2008; Markram et al., 2011). Current paradigms of fear conditioning include examples in which the CS and US stimuli do not overlap (Krabbe et al., 2019). Such a condition might seem to rule out the mechanisms in our paper. Nevertheless, the argument below suggests that the effects of the CS and US can cause an overlap in neuronal spiking of ECS, F, VIP, and SOM, even when CS and US inputs do not overlap.

Added in the revised version: Experimental recordings cannot speak to the rate of spiking of BLA neurons during US due to recording interference from the shock. However, evidence suggests that ECS activity should increase during the US due to the release of acetylcholine (ACh) from neurons in the basal forebrain (BF) (Rajebhosale et al., 2024b). Pyramidal cells of the BLA robustly express M1 muscarinic ACh receptors (McDonald and Mott, 2021). Thus, ACh from BF should elicit a depolarization in pyramidal cells. Indeed, the pairing of ACh with even low levels of spiking of BLA neurons results in a membrane depolarization that can last 7 – 10 s (Unal et al., 2015). Other modulators, including dopamine, may also play a role in producing the sustained activity. Activation of US leads to increased dopamine release in the BLA (Harmer and Phillips, 1999; Suzuki et al., 2002). D1 receptors are known to increase the membrane excitability of BLA projection neurons by lowering their spiking threshold (Kröner et al., 2005). Thus, neuromodulator release should induce higher spiking rates and more sustained activity in the ECS and F neurons during and after the presentation of US, thus ensuring a concomitant activation of ECS and fear (F) neurons necessary for STDP to take place. Thus, the activation of the US can lead to continued and higher firing rates of ECS and F. The effect of dopamine can last up to 20 minutes (Kröner et al., 2005). For CS-positive neurons, the ACh modulation coming from the firing of US may lead to a temporary extension of firing that is then amplified and continued by dopaminergic effects.

Added in the revised version: Hence, we suggest that a solution to the problem apparently posed by the non-overlap US and CS in some paradigms of auditory fear conditioning (Krabbe et al., 2019) may be solved by considering the roles of ACh and dopamine in the BLA. The model we have may be considered a “minimal” model that puts in by hand the overlap in activity due to the neuromodulation without explicitly modeling it. We have used the simplest way to model overlap without assumptions about timing specificity in the overlap. We note that, even though ECS and F neurons have the ability to fire continuously when ACh and dopamine are involved, the participation of the interneurons enforces periodic silence needed for the depression-dominated STDP.

Added in the revised version: Is STDP needed in fear conditioning?

Added in the revised version: The study in (Grewe et al., 2017) questions the validity of the Hebbian model in establishing associative learning during fear conditioning. There are several critiques we discuss here. The first critique is that Hebbian plasticity does not explain the experimental finding showing that both upregulation and downregulation of stimulus-evoked responses are present between coactive neurons. The upregulation is provided by our model, so the issue is the downregulation, which is not addressed by our model. However, our model highlights that coactivity alone does not create potentiation; the fine timing of the pre- and postsynaptic spikes determines whether there is potentiation or depression. Here, we find that PING networks are instrumental in setting up the fine timing for potentiation. We suggest that networks not connected to produce the PING may undergo depression when coactive.

Added in the revised version: The second critique raised by (Grewe et al., 2017) is that Hebbian plasticity alone does not explain why most of the cells exhibiting enhanced responses to the CS did not react to the US before fear conditioning. They suggest that neuromodulators may provide a third condition (besides the activity of the pre- and postsynaptic neurons) that changes the plasticity rule. Our model also does not explicitly address this experimental finding since it requires F to be initially activated by US in order for the fear association to be established. We agree that the fear cells described in (Grewe et al., 2017) may be depolarized by the US without reaching the spiking threshold; however, with neuromodulation provided during the fear training, the same input can lead to spiking, enabling the conditions for Hebbian plasticity. Our discussions above about how neuromodulators affect excitability are relevant to this point. We do not exclude that other forms of plasticity may play a role during fear conditioning in cells not initially activated by the US, but this is not the topic of our modeling study.

Added in the revised version: The third critique raised by (Grewe et al., 2017) is that Hebbian plasticity cannot explain why the majority of cells that were US- and CS-responsive before training have a reduced CS-evoked response afterward. The reduced response happens over multiple exposures of CS without US; this can involve processes similar to those present in fear extinction, which require plasticity in further networks, especially involving the infralimbic cortex (Milad and Quirk, 2002; Burgos-Robles et al., 2007). An extension of our model could investigate such mechanisms. In the fourth critique, (Grewe et al., 2017) suggests that the Hebbian plasticity rule cannot easily account for the reduction of the responses of many CS+-responsive cells, but not of the CS−-responsive cells. We suggest that the circuits involving paradigms similar to fear extinction do not involve the CS-cells.

Added in the revised version: Overall, we agree with (Grewe et al., 2017) that neuromodulators play a crucial role in fear conditioning, especially in prolonging the US- and CS-encoding activity as discussed in (see section “Assumptions and predictions of the model” in the Discussion), or even participating in changing the details of the plasticity rule. A possible follow-up of our work involves investigating how fear ensembles form and modify through fear conditioning and later stages. This follow-up work may involve using a tri-conditional rule, as suggested in (Grewe et al., 2017), in which the potential role of neuromodulators is taken into account in the plasticity rule in addition to the pre- and postsynaptic neuron activity. Another direction is to investigate a possible relationship between neuromodulation and a depression-dominated Hebbian rule.

Comparison with other models

Many computational models that study fear conditioning have been proposed in the last years; the list includes biophysically detailed models (e.g., (Li et al., 2009; Kim et al., 2013), firing rate models (e.g., (Krasne et al., 2011; Vlachos et al., 2011; Ball et al., 2012)), and connectionist models (e.g., (Edeline and Weinberger, 1992; Armony et al., 1997; Moustafa et al., 2013) (for a review see (Nair et al., 2016)). Both firing rate models and connectionist models use an abstract description of the interacting neurons or regions. The omission of biophysical details prevents such models from addressing questions concerning the roles of dynamics and biophysical details in fear conditioning, which is the aim of our model. There are also biophysically detailed models (Li et al., 2009; Kim et al., 2013, 2016; Feng et al., 2019), which differ from ours in both the physiology included in the model and the description of how plastic changes take place. One main difference in the physiology is that we differentiated among types of interneurons, since the fine timing produced for the latter was key to our use of rhythms to produce spike-timing-dependent plasticity. The origin of the gamma rhythm (but not the other rhythms) was investigated in (Feng et al., 2019), but none of these papers connected the rhythms to plasticity.

Earlier version: The most interesting difference between our work and that in (Li et al., 2009; Kim et al., 2013, 2016) is the modeling of plasticity. We use spike-timing-dependent plasticity rules. The models in (Li et al., 2009; Kim et al., 2013, 2016) were more mechanistic about how the plasticity takes place, starting with the known involvement of calcium with plasticity. Using a hypothesis about back propagation of spikes, the set of papers together come up with a theory that is consistent with STDP and other instantiations of plasticity (Shouval et al., 2002a, 2002b). For the purposes of our paper, this level of detail, though very interesting, was not necessary for our conclusions. By contrast, in order for the rhythms and the interneurons to have the dynamic roles they play in the model, we needed to restrict our STDP rule to ones that are depression-dominated. Our reading of (Shouval et al., 2002b) suggests to us that such subrules are possible outcomes of the general theory. Thus, there is no contradiction between the models, just a difference in focus; our focus was on the importance of the much-documented rhythms (Seidenbecher et al., 2003; Courtin et al., 2014a; Stujenske et al., 2014; Davis et al., 2017) in providing the correct spike timing. We showed in the Supplementary Information (“Classical Hebbian plasticity rule, unlike the depression-dominated one, shows potentiation even with no strict pre and postsynaptic spike timing”) that if the STDP rule was not depression dominated, the rhythms need not be necessary. We hypothesize that the necessity of strict timing enforced by the depression-dominated rule may foster the most appropriate association with fear at the expense of less relevant associations.

Revised version: The most interesting difference between our work and that in (Li et al., 2009; Kim et al., 2013, 2016) is the modeling of plasticity. We use spike-timing-dependent plasticity rules. The models in (Li et al., 2009; Kim et al., 2013, 2016) were more mechanistic about how the plasticity takes place, starting with the known involvement of calcium with plasticity. Using a hypothesis about back propagation of spikes, the set of papers together develop a theory that is consistent with STDP and other instantiations of plasticity (Shouval et al., 2002a, 2002b). For the purposes of our paper, this level of detail, though very interesting, was not necessary for our conclusions. By contrast, in order for the rhythms and the interneurons to have the dynamic roles they play in the model, we needed to restrict our STDP rule to ones that are depression-dominated. Our reading of (Shouval et al., 2002b) suggests to us that such subrules are possible outcomes of the general theory. Thus, there is no contradiction between the models, just a difference in focus; our focus was on the importance of the much-documented rhythms (Seidenbecher et al., 2003; Courtin et al., 2014a; Stujenske et al., 2014; Davis et al., 2017) in providing the correct spike timing. We showed in the Supplementary Information (“Classical Hebbian plasticity rule, unlike the depression-dominated one, shows potentiation even with no strict pre and postsynaptic spike timing”) that if the STDP rule was not depression dominated, the rhythms need not be necessary. We hypothesize that the necessity of strict timing enforced by the depression-dominated rule may foster the most appropriate association with fear at the expense of less relevant associations.

The dynamics of the VIP cell play a central role in the plasticity we investigate. This is in contrast to the cortical model in (Veit et al., 2023) for which VIP is essential for locally controlling gain and globally controlling coherence in gamma oscillations. In the model by Veit et al., 2023, the global control requires differences in long-range connectivity that are known to exist and are inserted in the model by hypothesis. Our paper shows how more detailed biophysics produces rhythms among the interneurons used in (Veit et al., 2023) and how these rhythms can produce the plasticity needed to construct those differences in long-range connectivity. Thus, though (Veit et al., 2023) shows that rhythms are not needed for some kinds of control once connectivity is established, our paper suggests that the same set of interneurons, with more detailed physiology, can support the establishment of appropriate connectivity as well as the control described in (Veit et al., 2023). We note that Veit et al., 2023 deal with cortical networks, while our model describes BLA networks; however, it is known that these networks are structurally related (Sah et al., 2003; Tovote et al., 2015; Polepalli et al., 2020).

Removed in the revised version: The study in (Grewe et al., 2017) suggests that associative learning resulting from fear conditioning induces both potentiation and depression among coactive excitatory neurons; coactivity was determined by calcium signaling and thus did not allow measurements of fine timing between spikes. In our model, we show how potentiation between coactive cells occurs when strict pre-post spike timing and appropriate pauses in the spiking activity arise. Depression happens when one or both of these components are not present. Thus, in our model, depression represents the absence of successful fear association and does not take part in the reshaping of the ensemble encoding the association, as instead suggested in (Grewe et al., 2017). A possible follow-up of our work involves investigating how fear ensembles form and modify through fear conditioning and later stages. This follow-up work may involve using a tri-conditional rule, as suggested in (Grewe et al., 2017), in which the potential role of neuromodulators is taken into account in addition to the pre- and postsynaptic neuron activity; this may lead to both potentiation and depression in establishing an associative memory.

Removed in the revised version: Where the rhythms originate, and by what mechanisms

Removed in the revised version: A recent experimental paper, (Antonoudiou et al., 2022), suggests that the BLA can intrinsically generate theta oscillations (3-12 Hz) detectable by LFP recordings under certain conditions, such as reduced inhibitory tone. They draw this conclusion in mice by removing the hippocampus, which can volume conduct to BLA, and noticing that other nearby brain structures did not display any oscillatory activity. Our model also supports the idea that intrinsic mechanisms in the BLA can support the generation of the low theta, high theta, and gamma rhythms.

Removed in the revised version: Although the BLA can produce these rhythms, this does not rule out that other brain structures also produce the same rhythms through different mechanisms, and these can be transmitted to the BLA. Specifically, it is known that the olfactory bulb produces and transmits the respiratory-related low theta (4 Hz) oscillations to the dorsomedial prefrontal cortex, where it organizes neural activity (Bagur et al., 2021). Thus, the respiratory-related low theta may be captured by BLA LFP because of volume conduction or through BLA extensive communications with the prefrontal cortex. Furthermore, high theta oscillations are known to be produced by the hippocampus during various brain functions and behavioral states, including during spatial exploration (Vanderwolf, 1969) and memory formation/retrieval (Raghavachari et al., 2001), which are both involved in fear conditioning. Similarly to the low theta rhythm, the hippocampal high theta can manifest in the BLA. It remains to understand how these other rhythms may interact with the ones described in our paper.

Removed in the revised version: Assumptions and predictions of the model

Removed in the revised version: The interneuron descriptions in the model were constrained by the electrophysiological properties reported in response to hyperpolarizing currents (Sosulina et al., 2010). Specifically, we modeled the three subtypes of VIP, SOM, and PV interneurons displaying bursting behavior, regular spiking with early spike-frequency adaptation, and regular spiking without spike-frequency adaptation, respectively. Focusing on VIP interneurons, we were able to model the bursting behavior by including the D-type potassium current. This current is thought to exist in the VIP interneurons in the cortex (Porter et al., 1998), but whether this current is also found in the BLA is still unknown. Similarly, we endowed the SOM interneurons with NaP- and H-currents, as the OLM cells in the hippocampus. Due to these currents, the VIP and SOM cells are able to show low- and high-theta oscillations, respectively. The presence of these currents and the neurons’ ability to exhibit oscillations in the theta range during fear conditioning and at baseline in BLA, which are assumptions of our model, should be tested experimentally.

Removed in the revised version: Our model, which is a first effort towards a biophysically detailed description of the BLA rhythms and their functions, does not include the neuron morphology, many other cell types, conductances, and connections that are known to exist in the BLA; models such as ours are often called “minimal models” and constitute the majority of biologically detailed models. Such minimal models are used to maximize the insight that can be gained by omitting details whose influence on the answers to the questions addressed in the model are believed not to be qualitatively important. We note that the absence of these omitted features constitutes hypotheses of the model: we hypothesize that the absence of these features does not materially affect the conclusions of the model about the questions we are investigating. Of course, such hypotheses can be refuted by further work showing the importance of some omitted features for these questions and may be critical for other questions. Our results hold when there is some degree of heterogeneity of cells of the same type, showing that homogeneity is not a necessary condition.

Removed in the revised version: Our study suggests that all the interneurons are necessary for associative learning provided that the STDP rule is depression-dominated. This prediction could be tested experimentally by selectively silencing each interneuron subtype in the BLA: if the associative learning is hampered by silencing any of the interneuron subtypes, this validates our study. Finally, the model prediction could be tested indirectly by acquiring more information about the plasticity rule involved in the BLA during associative learning. We found that all the interneurons are necessary to establish fear learning only in the case of a depression-dominated rule. This rule ensures that fine timing and pauses are always required for potentiation: interneurons provide both fine timing and pauses to pyramidal cells, making them crucial components of the fear circuit.

Removed in the revised version: The modeling of the interneurons assumes the involvement of various intrinsic currents; the inclusion of those currents can be considered hypotheses of the model. Our model predicts that blockade of D-current in VIP interneurons (or silencing VIP interneurons) will both diminish low theta and prevent fear learning. Finally, the model assumes the absence of significantly strong connections from the excitatory projection cells ECS to PV interneurons, unlike the ones from F to PV. Including those synapses would alter the PING rhythm created by the interactions between F and PV, which is crucial for fine timing between ECS and F needed for LTP.

Removed in the revised version: Finally, our model requires the effect of the CS and US inputs on the BLA neuron activity to overlap in time in order to instantiate fear learning. Despite paradigms involving both overlapping (delay conditioning, where US co-terminates with CS (e.g., (Lindquist et al., 2004)), or immediately follows CS (e.g., (Krabbe et al., 2019)) and non-overlapping (trace conditioning) CS/US inputs exist, we hypothesized that concomitant activity in CS- and US-encoding neuron activity should be crucial in both cases. This may be mediated by the memory effect due to metabotropic effects (Whittington et al., 1995) as suggested above, or by the contribution from other brain regions (see section “Involvement of other brain structures” in the Discussion). The fact that plasticity occurs with US memory trace is a consequence of our larger hypothesis that fear learning uses spike-timing-dependent plasticity; such a hypothesis about plasticity is common in the modeling literature.

Limitations and caveats

LFPs recorded in the experiments are thought to be mainly created by transmembrane currents in neurons located around the electrode and depend on several factors, including the morphology of the arborization of contributing neurons and the location of AMPA and GABA boutons (Katzner et al., 2009; Lindén et al., 2011; Łęski et al., 2013; Mazzoni et al., 2015). Since our model has no spatial extension, we used an LFP proxy; this proxy was shown to reflect the rhythmic output of the network, which we believe to be the essential result (for more details see Results “Increased low-theta frequency is a biomarker of fear learning”, and Supplementary Information “A higher low theta power increase emerges in LFP approximated with the sum of the absolute values of the currents compared to their linear sum”).

Earlier version: The use of small number of neurons raises the issue of model scalability. One way in which large models can be different from much smaller ones is in heterogeneity of the neurons of any given type. By using a network with a few neurons, we have begun the study of effects of heterogeneity. In general, the use of small numbers, which effectively assumes that each cell represents a synchronous subset of a larger population, replaces the use of gap-junctions that are known to exist in the cortex among VIP cells (Francavilla et al., 2018) as well as among SOM cells and among PV cells (Tremblay et al., 2016).

Revised version: The use of small number of neurons raises the issue of model scalability. One way in which large models can be different from much smaller ones is in heterogeneity of the neurons of any given type. By using a network with a few neurons of each type, we have begun the study of effects of heterogeneity. In general, the use of small numbers, which effectively assumes that each cell represents a synchronous subset of a larger population, replaces the use of gap-junctions that are known to exist in the cortex among VIP cells (Francavilla et al., 2018) as well as among SOM cells and among PV cells (Tremblay et al., 2016).

We do not explicitly model the biophysics of NMDA receptors. Rather, we model the effect of such receptors using the spike-timing-dependent plasticity resulting from such biophysics, as is commonly done when modeling STDP (Song et al., 2000). Also, our neurons are single-compartment, so do not build in the spatial structure known to exist on the dendrites (Blair et al., 2001; Bennett et al., 2019).

Our model assumes that initial stages of fear learning can be accomplished entirely within the amygdala, though it is known that other structures in the brain are important for modulating networks related to fear. Much of the work involving the prefrontal cortex and the hippocampus relates to fear extinction, which is not addressed in this paper.

Added in the revised version: Summary

Removed in the revised version: Summary and significance

We have shown how networks of amygdala neurons, including multiple types of interneurons, can work together to produce plasticity needed for fear learning. The coordination necessary to produce the plasticity requires the involvement of multiple rhythms. Thus, our paper both accounts for the experimental evidence showing such amygdala rhythms exist and points to their central role in the mechanisms of plasticity involved in associative learning. These mechanisms may be common to other types of associative learning, as similar interneuron subtypes and connectivity are ubiquitous in the cortex (Sah et al., 2003; Tovote et al., 2015; Polepalli et al., 2020).

Removed in the revised version: Supporting information